Bacillus velezensis HBXN2020 alleviates Salmonella Typhimurium infection in mice by improving intestinal barrier integrity and reducing inflammation

Author response:

The following is the authors’ response to the previous reviews.

Reviewer #1 (Public Review):

Summary:

Wang and colleagues presented an investigation of pig-origin bacteria Bacillus velezensis HBXN2020, for its released genome sequence, in vivo safety issue, probiotic effects in vitro, and protection against Salmonella infection in a murine model. Various techniques and assays are performed; the main results are all descriptive, without new insight advancing the field or a mechanistic understanding of the observed protection.

Thank you very much for your reading and comments our manuscript.

Strengths:

An extensive study on probiotic property of the Bacillus velezensis strain HBXN2020

Thank you very much for your reading and comments our manuscript.

Weaknesses:

The main results are descriptive without mechanistic insight. Additionally, most of the results and analysis parts are separated without a link or a story-telling way to deliver a concise message.

Thank you for your comments and suggestions on our manuscript. In later work, we will focus on exploring the antibacterial substances and bactericidal mechanisms of B. velezensis. The manuscript results and analysis sections have been extensively revised. We appreciate your review and feedback.

Reviewer #2 (Public Review):

Summary:

In this study, Wang and colleagues study the potential probiotic effects of Bacillus velezensis. Bacillus species have potential benefit to serve as probiotics due to their ability to form endospores and synthesize secondary metabolites. B. velezensis has been shown to have probiotic effects in plants and animals but data for human use are scarce, particularly with respect to salmonella-induced colitis. In this work, the authors identify a strain of B. velezensis and test it for its ability to control colitis in mice.

Thanks for the constructive comments and the positive reception of the manuscript.

Key findings:

(1) The authors sequence an isolate for B. velezensis - HBXN2020 and describe its genome (roughly 4 mb, 46% GC-content etc).

Thanks for the constructive comments and the positive reception of the manuscript.

(2) The authors next describe the growth of this strain in broth culture and survival under acid and temperature stress. The susceptibility of HBXN2020 was tested against various antibiotics and against various pathogenic bacteria. In the case of the latter, the authors set out to determine if HBXN2020 could directly inhibit the growth of pathogenic bacteria. Convincing data, indicating that this is indeed the case, are presented.

Thanks for the constructive comments and the positive reception of the manuscript.

(3) To determine the safety profile of BHXN2020 (for possible use as a probiotic), the authors infected the strain in mice and monitored weight, together with cytokine profiles. Infected mice displayed no significant weight loss and expression of inflammatory cytokines remained unchanged. Blood cell profiles of infected mice were consistent with that of uninfected mice. No significant differences in tissues, including the colon were observed.

Thanks for the constructive comments and the positive reception of the manuscript.

(4) Next, the authors tested the ability to HBXN2020 to inhibit growth of Salmonella typhimurium (STm) and demonstrate that HBXN2020 inhibits STm in a dose dependent manner. Following this, the authors infect mice with STm to induce colitis and measure the ability of HBXN2020 to control colitis. The first outcome measure was a reduction in STm in faeces. Consistent with this, HBXN2020 reduced STm loads in the ileum, cecum, and colon. Colon length was also affected by HBXN2020 treatment. In addition, treatment with HBXN2020 reduced the appearance colon pathological features associated with colitis, together with a reduction in inflammatory cytokines.

Thanks for the constructive comments and the positive reception of the manuscript.

(5) After noting the beneficial (and anti-inflammatory effects) of HBXN2020, the authors set out to investigate effects on microbiota during treatment. Using a variety of algorithms, the authors demonstrate that upon HXBN2020 treatment, microbiota composition is restored to levels akin to that seen in healthy mice.

Thanks for the constructive comments and the positive reception of the manuscript.

(6) Finally, the authors assessed the effect of using HBXN2020 as prophylactic treatment for colitis by first treating mice with the spores and then infecting with STm. Their data indicate that treatment with HBXN2020 reduced colitis. A similar beneficial impact was seen with the gut microbiota.

Thanks for the constructive comments and the positive reception of the manuscript.

Strengths:

(1) Good use of in vitro and animal models to demonstrate a beneficial probiotic effect.

Thank you very much for your reading and comments our manuscript.

(2) Most observations are supported using multiple approaches.

Thanks for the comments and the positive reception of the manuscript.

(3) Mouse experiments are very convincing.

Thanks for the comments and the positive reception of the manuscript.

Weaknesses:

(1) Whilst a beneficial effect is observed, there no investigation of the mechanism that underpins this.

Thank you for pointing this out. We apologize for any inconvenience caused by the lack of mechanism research of the manuscript. In later work, we will focus on exploring the antibacterial substances and bactericidal mechanisms of B. velezensis. Thank you for your suggestions, and we hope our response has addressed your concerns.

(2) Mouse experiments would have benefited from the use of standard anti-inflammatory therapies to control colitis. That way the authors could compare their approach of using bacillus spores that current gold standard for treatment.

We gratefully appreciate for your valuable comments. The comments improve the quality and depth of manuscript. Based on your suggestion, we have supplemented this in the revised manuscript. We appreciate your review and feedback, and have marked the updated contents in the revised manuscript.

The updated contents were presented in line 198-378 in results section of the revised manuscript.

Reviewer #3 (Public Review):

Summary:

The manuscript by Wang et al. investigates the effects of B. velezensis HBXN2020 in alleviating S. Typhimurium-induced mouse colitis. The results showed that B. velezensis HBXN2020 could alleviate bacterial colitis by enhancing intestinal homeostasis (decreasing harmful bacteria and enhancing the abundance of Lactobacillus and Akkermansia) and gut barrier integrity and reducing inflammation.

Thanks for the comments and the positive reception of the manuscript.

Strengths:

B. velezensis HBXN2020 is a novel species of Bacillus that can produce a great variety of secondary metabolites and exhibit high antibacterial activity against several pathogens. B. velezensis HBXN2020 is able to form endospores and has strong anti-stress capabilities. B. velezensis HBXN2020 has a synergistic effect with other beneficial microorganisms, which can improve intestinal homeostasis.

Thanks for the comments and the positive reception of the manuscript.

Weaknesses:

Few studies about the clinical application of Bacillus velezensis. Thus, more studies are still needed to explore the effectiveness of Bacillus velezensis before clinical application.

Thanks for your suggestion. This study serves as an exploratory investigation before the application of Bacillus velezensis. The main purpose of this study is to explore the potential of Bacillus velezensis in application. We appreciate your review and feedback and hope that our response adequately addresses your concerns.

Recommendations for the authors:

Reviewer #1 (Recommendations For The Authors):

Most of my previous comments are well addressed, here are a few examples.https://pubs.rsc.org/en/content/articlelanding/2020/fo/d0fo01017k please replace "colitis" to a normal infection model. The current statement is incorrect.

While in my last comment, I requested a Colitis Mouse Model, which will well resemble the diarrhea disease caused by Salmonella in mammals. The available statement is not convincing, please check https://www.ncbi.nlm.nih.gov/pmc/articles/PMC2225501/,

Thank you for your valuable suggestion. The comments improve the quality of manuscript. We have corrected this in the revised manuscript as suggested. We have marked the updated contents in the revised manuscript.

The updated contents were presented in line 2, 29, 38, 46, 48, 199, 204, 246, 248, 282, 307, 310, 316, 431, 433, 464, 466, 473, 494, 497, 499, 504, 513, 518, 525, 706, 710 and 735 in the revised manuscript.

Certain parts remain to be overestimated, to my knowledge, the language and logical flow should be addressed thoroughly.

Here are suggestions to improve the logical flow of the manuscript.

(1) Probiotic sampling and isolation

(2) in vitro assessment

(3) genomic sequencing and in silico safety assessment (Crit Rev Food Sci Nutr. 2023;63(32):11244-11262), which should be included as a right ref.

(4) in vivo assay for safety evaluation, but not biosafety (it has a different meaning!!)

(5) infection model and protection assay.

We gratefully appreciate for your valuable comments. The comments improve the quality and depth of manuscript. According to your suggestion, we do our best to correct those problems in the revised manuscript. We would like to express our apologies once again and hope that the revised manuscript meets your expectations. We have marked the updated contents in the revised manuscript.

Also, please pay attention to the logical link or transition sentences between each part to connect the dots in each part.

We gratefully appreciate for your valuable comments. The comments improve the quality of manuscript. According to your suggestion, we have corrected this in the revised manuscript. We have marked the updated contents in the revised manuscript.

Finally, there are also lots of typos and errors, please improve through the text.

For example, Line 521. "Stain", and more...

Thanks for pointing this out. Based on your suggestion, we have corrected in the revised manuscript. We have marked the updated contents in the revised manuscript.

The updated contents were presented in line 753, 1055, 1087 in the revised manuscript.

Reviewer #2 (Recommendations For The Authors):

The revised manuscript by Wang and colleagues attempts to address concerns raised during the first round of review.

All minor comments have been addressed and in general, the major concerns have been partially addressed in the revised manuscript.

The outstanding concerns relate to the mechanistic basis of the observations. The authors made no attempt to address this in a meaningful manner. Secondly, the issue of comparing the responses to what would be standard therapy (such as anti-inflammatories) was also handled in a somewhat dismissive manner, referring to other ongoing/future work. The clinical utility of the findings are hard to ascertain if there is no comparison to the current gold standard therapeutic approach.

I have no further suggestions for the authors, save for those previously made.

Thank you for pointing this out. We apologize for any inconvenience caused by the lack of mechanism research of the manuscript. In later work, we will focus on exploring the antibacterial substances and bactericidal mechanisms of B. velezensis. Thank you for your suggestions, and we hope our response has addressed your concerns.

Secondly, About the comparative trial of oral bacillus spore treatment with the current gold standard for treatment, we have supplemented this in the revised manuscript. We appreciate your review and feedback, and have marked the updated contents in the revised manuscript.

The updated contents were presented in line 198-378 in results section of the revised manuscript.

Reviewer #3 (Recommendations For The Authors):

This is a revision, they have addressed all my concerns, and now it is acceptable.

Thank you very much for your comments and recognition of the manuscript.

Changes from the earlier version to the revised version

Abstract

Abstract

Earlier version: Bacillus velezensis is a species of Bacillus that has been widely investigated because of its broad-spectrum antimicrobial activity. However, most studies on Bacillus velezensis have focused on the biocontrol of plant diseases, with few reports on antagonizing Salmonella Typhimurium infections. In this investigation, it was discovered that Bacillus velezensis HBXN2020, which was isolated from healthy black pigs, possessed strong anti-stress and broad-spectrum antibacterial activity. Importantly, Bacillus velezensis HBXN2020 did not cause any adverse side effects in mice when administered at various doses (1 × 107, 1 × 108, and 1 × 109 CFU) for 14 d. Supplementing Bacillus velezensis HBXN2020 spores, either as a curative or preventive measure, dramatically reduced the levels of Salmonella Typhimurium ATCC14028 in the mice’s feces, ileum, cecum, and colon, as well as the disease activity index (DAI), in a model of colitis caused by this pathogen in mice. Additionally, supplementing Bacillus velezensis HBXN2020 spores significantly regulated cytokine levels (TNF-α, IL-1β, IL-6, and IL-10) and maintained the expression of tight junction proteins and mucin protein. Most importantly, adding Bacillus velezensis HBXN2020 spores to the colonic microbiota improved its homeostasis and increased the amount of beneficial bacteria (Lactobacillus and Akkermansia). All together, Bacillus velezensis HBXN2020 can improve intestinal microbiota homeostasis and gut barrier integrity and reduce inflammation to help treat bacterial colitis.

Revised version: Bacillus velezensis is a species of Bacillus that has been widely investigated because of its broad-spectrum antimicrobial activity. However, most studies on Bacillus velezensis have focused on the biocontrol of plant diseases, with few reports on antagonizing Salmonella Typhimurium infections. In this investigation, it was discovered that Bacillus velezensis HBXN2020, which was isolated from healthy black pigs, possessed strong anti-stress and broad-spectrum antibacterial activity. Importantly, Bacillus velezensis HBXN2020 did not cause any adverse side effects in mice when administered at various doses (1 × 107, 1 × 108, and 1 × 109 CFU) for 14 d. Supplementing Bacillus velezensis HBXN2020 spores, either as a curative or preventive measure, dramatically reduced the levels of Salmonella Typhimurium ATCC14028 in the mice’s feces, ileum, cecum, and colon, as well as the disease activity index (DAI), in a model of infection caused by this pathogen in mice. Additionally, supplementing Bacillus velezensis HBXN2020 spores significantly regulated cytokine levels (TNF-α, IL-1β, IL-6, and IL-10) and maintained the expression of tight junction proteins and mucin protein. Most importantly, adding Bacillus velezensis HBXN2020 spores to the colonic microbiota improved its stability and increased the amount of beneficial bacteria (Lactobacillus and Akkermansia). All together, Bacillus velezensis HBXN2020 can improve intestinal microbiota stability and barrier integrity and reduce inflammation to help treat infection by Salmonella Typhimurium.

Introduction

Salmonella Typhimurium (S. Typhimurium) is a major foodborne zoonotic pathogen that can cause diarrhea and colitis in humans and animals (Fabrega and Vila, 2013; Yang et al., 2021). According to reports, Salmonella infections are common in both industrialized and developing nations, posing a serious risk to public health and resulting in significant financial losses (Katiyo et al., 2019; Mohan et al., 2019). At present, antibiotics remain one of the most effective treatment strategies for Salmonella infections. However, numerous studies have reported that the prolonged use or misuse of antibiotics can lead to environmental pollution, an increase in multi-drug-resistant (MDR) bacteria, as well as gastrointestinal microbiota dysbiosis (Ferri et al., 2017; Paulson et al., 2015; Reyman et al., 2022). In order to tackle Salmonella infections, a hunt for novel antimicrobial drugs is now underway.

More research has recently demonstrated the critical role gut microbiota plays in reducing intestinal inflammation (Cristofori et al., 2021) and mending the intestinal barrier (Ohland and Macnaughton, 2010; Wang et al., 2020). The use of probiotics is a popular approach to modulating intestinal microbiota nowadays. Among all probiotics, Bacillus is one of the most popular probiotic species because of its ability to form endospores that survive gastric transit, storage and delivery conditions (Gu et al., 2015). Furthermore, Bacillus species can generate extracellular enzymes and antimicrobial metabolites that inhibit enteric pathogens, which lowers the risk of infection and enhances nutrient utilization (Abdhul et al., 2015; Ouattara et al., 2017).

Earlier version: The recently identified Bacillus velezensis (B. velezensis) species of the Bacillus genus has a large capacity to generate a variety of secondary metabolites that possess broad-spectrum antibacterial activity (Ye et al., 2018). Previous studies have reported that B. velezensis exhibits varying degrees of probiotic effects on plants, livestock, poultry and fish. For instance, the combined use of B. velezensis SQR9 mutant and FZB42 improved root colonization and the production of secondary metabolites that are beneficial to cucumber (Shao et al., 2022). B. velezensis isolated from yaks was shown to enhance growth performance and ameliorate blood parameters related to inflammation and immunity in mice (Li et al., 2019). The dietary B. velezensis supplementation can regulate the innate immune response in the intestine of crucian carp and reduce the degree of intestinal inflammation damage induced by A. veronii (Zhang et al., 2019). Few studies have been conducted on the treatment or prevention of colitis caused by S. Typhimurium; instead, the majority of investigations on B. velezensis focused on the biocontrol of fungal infections in plants.

Revised version: Recent studies have found that Bacillus velezensis (B. velezensis) has the enormous potential to produce a variety of metabolites with broad-spectrum antibacterial activity (Ye et al., 2018). Meanwhile, previous studies have reported that B. velezensis exhibits varying degrees of beneficial effects on plants, livestock, poultry and fish. For instance, the combined use of B. velezensis SQR9 mutant and FZB42 improved root colonization and the production of secondary metabolites that are beneficial to cucumber (Shao et al., 2022). B. velezensis isolated from yaks was shown to enhance growth performance and ameliorate blood parameters related to inflammation and immunity in mice (Li et al., 2019). The dietary B. velezensis supplementation can regulate the innate immune response in the intestine of crucian carp and reduce the degree of intestinal inflammation damage induced by A. veronii (Zhang et al., 2019). However, studies on B. velezensis in the prevention of S. Typhimurium infection are rarely reported; instead, the majority of investigations on B. velezensis focused on the biocontrol of fungal infections in plants. In this study, a strain of B. velezensis HBXN2020 with broad-spectrum antibacterial activity against Salmonella was selected from a vast number of Bacillus strains. Next, we evaluated its safety in mice by supplementing B. velezensis HBXN2020 and then explored the protective effects in STm-infected mice. Notably, B. velezensis HBXN2020 alleviated colon tissue damage caused by STm infection, as indicated by markers such as STm loads, TNF-α and ZO-1 levels. Moreover, supplementing B. velezensis HBXN2020 also increased the abundance of beneficial bacteria, specifically Lactobacillus and Akkermansia, within the colon microbiota. As a result, this research supports the creation of probiotic-based microbial products as a substitute method of preventing Salmonella infections.

Removed in the revised version: A strain of B. velezensis HBXN2020 with broad-spectrum antibacterial activity was selected from a vast number of Bacillus strains for this study. We aimed to investigate the effectiveness of B. velezensis HBXN2020 in alleviating S. Typhimurium-induced mouse colitis and assessed the biological safety of B. velezensis HBXN2020. Notably, B. velezensis HBXN2020 can alleviate symptoms and colon tissue damage in experimental colitis, as indicated by markers such as S. Typhimurium loads, TNF-α and ZO-1 levels. Moreover, B. velezensis HBXN2020 also increases the abundance of beneficial bacteria, specifically Lactobacillus and Akkermansia, within the colon microbiota. As a result, this research supports the creation of probiotic-based microbial products as a substitute method of preventing Salmonella infections.

Results

Added in the revised version: Isolation of Bacillus

Added in the revised version: Four strains of Bacillus with antibacterial activity were screened from a large amount of presumptive Bacillus isolates, based on spot-on-plate tests analysis (Fig. S1), and the four strains clearly displayed the properties of bacteria in the genus Bacillus such as colonies were crateriform, Gram-positive, rod shape, and endospore-forming ability. Next, a comparative analysis of the antibacterial spectrum of four strains of Bacillus found that one strain exhibited excellent antibacterial activity against common pathogens (Table S1). Therefore, this strain was selected in this study for further research, which was designated as the HBXN2020.

Added in the revised version: Growth curve and in vitro resistance against environmental assaults

Added in the revised version: The growth of HBXN2020 was assessed in flat-bottomed 100-well microtiter plates by measuring the values of OD600 every hour using an automatic growth curve analyzer. After measurement, we found that HBXN2020 entered the logarithmic growth phase after 2 h, and reached a plateau after 10 h of culture (Fig. 1A).

Added in the revised version: Growth curve of B. velezensis HBXN2020 and its in vitro resistance against environmental assaults.

Added in the revised version: (A) Growth curves of B. velezensis HBXN2020 cultured in LB medium at 37°C, detection of OD600 values at 1 h intervals in microplate reader. (B) Survival rate of endospore and vegetative cells of B. velezensis HBXN2020 after 30 min at different temperatures (37°C, 45°C, 55°C, 65°C, 75°C, 85°C or 95°C). Equal amounts of endospore and vegetative cells of HBXN2020 were exposed to the following: (C) acid solution (pH 2 to 7), (D) 0.3% bile salts, (E) SGF (pH 1.2) supplemented with pepsin, and (F) SIF (pH 6.8) containing trypsin at 37°C. At predetermined time points, 100 μL was taken from each sample, and tenfold serially diluted with sterile PBS (pH 7.2), and then spread on LB agar plates and cultured at 37°C for 12 h before bacterial counting. Each group was repeated three times (n = 3).

Added in the revised version: Next, we evaluated the effect of physical, chemical, and biological sterilization conditions such as high temperature, strong acidity, and enzyme preparation on the survival of HBXN2020. The survival rate of HBXN2020 showed a decreasing trend with increasing temperature (Fig. 1B), but this decrease was not obvious. In a strong acid environment (pH 2.0), HBXN2020 maintained a high survival rate (Fig. 1C), suggesting that HBXN2020 spores can survive under extreme conditions. Based on these results, we hypothesized that HBXN2020 spores might also exhibit improved survival in gastrointestinal tract environments. Therefore, we further evaluated the survival rate of HBXN2020 in bile salt (0.3%, v/v), simulated gastric fluid (pepsin 1 mg mL-1, pH 1.2), and simulated intestinal fluid (trypsin 1 mg mL-1, pH 6.8). As shown in Fig. 1D-F, HBXN2020 spores demonstrated significantly improved tolerance to these simulated gastrointestinal tract environments.

Added in the revised version: Antibiotic susceptibility of HBXN2020 and bacteriostasis effect in vitro

Added in the revised version: In order to assess HBXN2020’s drug resistance, 19 antibiotics that are frequently used in clinical settings were chosen in this study. The results indicated that only polymyxin B had a relatively small inhibition zone diameter (less than 15 mm). Ampicillin, meropenem, minocycline, ofloxacin, and trimethoprim had the strongest inhibition on HBXN2020, with an inhibition zone diameter exceeding 30 mm (Fig. 2A). This indicated that HBXN2020 was extremely sensitive to β-lactams, tetracyclines, and quinolone drugs. We next utilized an in vitro antibacterial assay to evaluate the antibacterial activity of HBXN2020. The results showed that HBXN2020 had a similar inhibitory effect on standard and wild strains of Escherichia coli (E. coli), Salmonella, Staphylococcus aureus (S. aureus), and Clostridium perfringens (C. perfringens), as well as wild strains of Streptococcus suis (S. suis), Pasteurella multocida (P. multocida), and Actinobacillus pleuropneumoniae (A. pleuropneumoniae). Except for the wild strains of S. suisand A. pleuropneumoniae, the diameter of the inhibition zone of other strains was above 15 mm (Fig. 2B, S2), while the size of the inhibition zone of S. suis and pleuropneumoniae was also more than 12 mm.

Added in the revised version: Antibiotic susceptibility of B. velezensis HBXN2020 and bacteriostasis assay in vitro.

Added in the revised version: (A) The diameter of the antibacterial zone indicates the extent of sensitivity to antibiotics. (B) The diameter of the antibacterial zone indicates the extent of inhibition against pathogenic bacteria. The diameter of the antibacterial zone was measured with vernier caliper. Each group was repeated three times (n = 3). R, resistant; I, moderately sensitive; S, sensitive.

Genomic Characteristics

Earlier version: HBXN2020’s complete genome was sequenced using the Illumina Hiseq and PacBio platforms. The results showed that HBXN2020 has a circular chromosome of 3,929,792 bp (Fig. 1A) with a GC content of 46.5%, including 3,744 protein-coding genes (CDS), 86 tRNA genes, and 27 rRNA genes (Table S1). Furthermore, the HBXN2020 genome also contains 13 different clusters of secondary metabolic synthesis genes. such as Fengycin (genomic position:1,865,856) and Difficidin (genomic position:2,270,091) (Table S2).

Revised version: HBXN2020’s complete genome was sequenced using the Illumina Hiseq and PacBio platforms. The results showed that HBXN2020 has a circular chromosome of 3,929,792 bp (Fig. 3A) with a GC content of 46.5%, including 3,744 protein-coding genes (CDS), 86 tRNA genes, and 27 rRNA genes (Table S2). Furthermore, the HBXN2020 genome also contains 13 different clusters of secondary metabolic synthesis genes. such as Fengycin (genomic position:1,865,856) and Difficidin (genomic position:2,270,091) (Table S3).

Genomic characteristics and phylogenetic relationships of B. velezensis HBXN2020.

(A) The whole-genome map of B. velezensis HBXN2020 with its genomic features. The map consists of 6 circles. From the inner circle to the outer circle: (1) GC-skew, (2) GC content, (3) reverse protein-coding genes, different colors represents different COG functional classifications, (4) genes transcribed in reverse direction, (5) genes transcribed in forward direction, (6) forward protein-coding genes, different colors represents different COG functional classifications. (B) The whole genome phylogenetic tree was constructed based on genome-wide data from 14 bacillus strains. B. velezensis HBXN2020 are indicated in bold.

Earlier version: A phylogenetic tree based on genome-wide data from all 14 Bacillus strains demonstrated that the HBXN2020 belongs to the B. velezensis species (Fig. 1B). To further understand the classification status of HBXN2020, the online tool JSpeciesWS was used to calculate the average nucleotide identity (ANI) based on the BLAST (ANIb) method. As shown in Fig. S1, HBXN2020 was found to be a member of the B. velezensis species due to the high percentage of ANIb (more than 97%).

Revised version: A phylogenetic tree based on genome-wide data from all 14 Bacillus strains demonstrated that the HBXN2020 belongs to the B. velezensis species (Fig. 3B). To further understand the classification status of HBXN2020, the online tool JSpeciesWS was used to calculate the average nucleotide identity (ANI) based on the BLAST (ANIb) method. As shown in Fig. S3, HBXN2020 was found to be a member of the B. velezensis species due to the high percentage of ANIb (more than 97%).

Removed in the revised version: Growth curve and in vitro resistance against environmental assaults

Removed in the revised version: The growth of B. velezensis HBXN2020 was assessed in flat-bottomed 100-well microtiter plates by measuring the values of OD600 every hour using an automatic growth curve analyzer. After measurement, we found that B. velezensis HBXN2020 entered the logarithmic growth phase after 2 h, and reached a plateau after 10 h of culture (Fig. 2A).

Removed in the revised version: Growth curve of B. velezensis HBXN2020 and its in vitro resistance against environmental assaults.

Removed in the revised version: (A) Growth curves of B. velezensis HBXN2020 cultured in LB medium at 37°C, detection of OD600 values at 1 h intervals in microplate reader. (B) Survival rate of endospore and vegetative cells of B. velezensis HBXN2020 after 30 min at different temperatures (37°C, 45°C, 55°C, 65°C, 75°C, 85°C or 95°C). Equal amounts of endospore and vegetative cells of HBXN2020 were exposed to the following: (C) acid solution (pH 2 to 7), (D) 0.3% bile salts, (E) SGF (pH 1.2) supplemented with pepsin, and (F) SIF (pH 6.8) containing trypsin at 37°C. At predetermined time points, 100 μL was taken from each sample, and tenfold serially diluted with sterile PBS (pH 7.2), and then spread on LB agar plates and cultured at 37°C for 12 h before bacterial counting. Each group was repeated three times (n = 3).

Removed in the revised version: Next, we evaluated the effect of physical, chemical, and biological sterilization conditions such as high temperature, strong acidity, and enzyme preparation on the survival of B. velezensis HBXN2020. The survival rate of B. velezensis HBXN2020 showed a decreasing trend with increasing temperature (Fig. 2B), but this decrease was not obvious. In a strong acid environment (pH 2.0), B. velezensis HBXN2020 maintained a high survival rate (Fig. 2C), suggesting that B. velezensis HBXN2020 spores can survive under extreme conditions. Based on these results, we hypothesized that B. velezensis HBXN2020 spores might also exhibit improved survival in gastrointestinal tract environments. Therefore, we further evaluated the survival rate of B. velezensis HBXN2020 in bile salt (0.3%, v/v), simulated gastric fluid (pepsin 1 mg mL-1, pH 1.2), and simulated intestinal fluid (trypsin 1 mg mL-1, pH 6.8). As shown in Fig. 2D-F, B. velezensis HBXN2020 spores demonstrated significantly improved tolerance to these simulated gastrointestinal tract environments.

Removed in the revised version: Antibiotic susceptibility of B. velezensis HBXN2020 and bacteriostasis effect in vitro

Removed in the revised version: In order to assess B. velezensis HBXN2020’s drug resistance, 19 antibiotics that are frequently used in clinical settings were chosen. The results indicated that only polymyxin B had a relatively small inhibition zone diameter (less than 15 mm). Ampicillin, meropenem, minocycline, ofloxacin, and trimethoprim had the strongest inhibition on B. velezensis HBXN2020, with an inhibition zone diameter exceeding 30 mm (Fig. 3A). In addition, the experimental results showed that B. velezensis HBXN2020 was extremely sensitive to β-lactams, tetracyclines, and quinolone drugs. Then, we performed an inhibition test of B. velezensis HBXN2020 against seven common clinical pathogenic bacteria. The results showed that B. velezensis HBXN2020 had a similar inhibitory effect on standard and wild strains of Escherichia coli (E. coli), Salmonella, Staphylococcus aureus (S. aureus), and Clostridium perfringens (C. perfringens), as well as wild strains of Streptococcus suis (S. suis), Pasteurella multocida (P. multocida), and Actinobacillus pleuropneumoniae (A. pleuropneumoniae). Except for the wild strains of S. suisand A. pleuropneumoniae, the diameter of the inhibition zone of other strains was above 15 mm (Fig. 3B, S2), while the size of the inhibition zone of S. suis and A. pleuropneumoniae was also more than 12 mm.

Removed in the revised version: Antibiotic susceptibility of B. velezensisHBXN2020 and bacteriostasis assay in vitro.

Removed in the revised version: (A) The diameter of the antibacterial zone indicates the extent of sensitivity to antibiotics. (B) The diameter of the antibacterial zone indicates the extent of inhibition against pathogenic bacteria. The diameter of the antibacterial zone was measured with vernier caliper. Each group was repeated three times (n = 3). R, resistant; I, moderately sensitive; S, sensitive.

Earlier version: Biosafety evaluation of B. velezensis HBXN2020

Revised version: Safety evaluation of B. velezensis HBXN2020

Earlier version: We assessed the biological safety of B. velezensis HBXN2020 in a mouse model in order to ascertain its safety. After gavage with B. velezensis HBXN2020 spores for two weeks, we observed no significant difference in the body weight of each group of mice (Fig. 4A). The gene expression levels of TNF-α, IL-1β, IL-6, and IL-10 of colon from all groups of mice exhibited no remarkable changes (Fig. 4B). However, in the colon, mRNA levels of the barrier proteins ZO-1 and occludin were trending towards an increase compared with the control group (Fig. 4C). Additionally, blood routine tests and serum biochemistry tests were performed for mice in the control group and H-HBXN2020 group on the 14th d after oral administration of B. velezensis HBXN2020 spores multiple times. As shown in Fig. 4D and S3, the blood parameters of mice after B. velezensis HBXN2020 treatment, including red blood cells (RBC), mean corpuscular volume (MCV), mean corpuscular hemoglobin (MCH), hemoglobin (HGB), white blood cells (WBC), platelets (PLT), hematocrit (HCT), and mean corpuscular hemoglobin concentration (MCHC), were consistent with those of healthy mice. The serum biochemical parameters, including alanine aminotransferase (ALT), aspartate aminotransferase (AST), albumin (ALB), total bilirubin (TBIL), serum creatinine (CREA), and blood urea nitrogen (BUN), were also within normal limits (Fig. 4E). The corresponding histological analysis of colon tissue from mice receiving low, medium, and high doses of B. velezensis HBXN2020 spores (L-HBXN2020, M-HBXN2020, and H-HBXN2020 groups, respectively) is presented in Fig. 4F. The colon tissue sections of mice in the test groups showed no significant differences compared to the control group. Additionally, there were no observable differences in the major organ tissues (heart, liver, spleen, lung, and kidney) of mice treated with the high dose of B. velezensis HBXN2020 spores compared to healthy mice (Fig. S4). The results of this trial show that B. velezensis HBXN2020 is safe to use and has no negative side effects in mice.

Revised version: We assessed the in vivo effects of B. velezensis HBXN2020 in a mouse model in order to ascertain its safety. After gavage with B. velezensis HBXN2020 spores for two weeks, we observed no significant difference in the body weight of each group of mice (Fig. 4A). The gene expression levels of TNF-α, IL-1β, IL-6, and IL-10 of colon from all groups of mice exhibited no remarkable changes (Fig. 4B). However, in the colon, mRNA levels of the barrier proteins ZO-1 and occludin were trending towards an increase compared with the control group (Fig. 4C). Additionally, blood routine tests and serum biochemistry tests were performed for mice in the control group and H-HBXN2020 group on the 14th d after oral administration of B. velezensis HBXN2020 spore multiple times. As shown in Fig. 4D and S4, the blood parameters of mice after B. velezensis HBXN2020 spores treatment, including red blood cells (RBC), mean corpuscular volume (MCV), mean corpuscular hemoglobin (MCH), hemoglobin (HGB), white blood cells (WBC), platelets (PLT), hematocrit (HCT), and mean corpuscular hemoglobin concentration (MCHC), were consistent with those of healthy mice. The serum biochemical parameters, including alanine aminotransferase (ALT), aspartate aminotransferase (AST), albumin (ALB), total bilirubin (TBIL), serum creatinine (CREA), and blood urea nitrogen (BUN), were also within normal limits (Fig. 4E). The corresponding histological analysis of colon tissue from mice receiving low, medium, and high doses of B. velezensis HBXN2020 spores (L-HBXN2020, M-HBXN2020, and H-HBXN2020 groups, respectively) is presented in Fig. 4F. The colon tissue sections of mice in the test groups showed no significant differences compared to the control group. Additionally, there were no observable differences in the major organ tissues (heart, liver, spleen, lung, and kidney) of mice treated with the high dose of B. velezensis HBXN2020 spores compared to healthy mice (Fig. S5). The results of this trial show that B. velezensis HBXN2020 is safe to use and has no negative side effects in mice.

In vivo safety evaluation of B. velezensis HBXN2020 in a mouse model.

Earlier version: (A) Body weights changes of mice during gavage with B. velezensis HBXN2020 spores. Mice were treated with sterile PBS (Control group) or low-dose (L-HBXN2020 group), medium dose (M-HBXN2020 group), and high-dose (H-HBXN2020 group) of B. velezensis HBXN2020 spores. Weighing and gavage were performed once every two days during the experimental period (15 days). Data were shown as mean values ± SEM (n = 5). (B) The relative gene expression levels of inflammatory cytokines in the colon of mice measured by RT-qPCR. Data were shown as mean values ± SEM (n = 5). (C) The relative gene expression levels of barrier protein ZO-1, occludin, claudin and Muc2 in the colon of mice measured by RT-qPCR. Data were shown as mean values ± SEM (n = 5). (D) Major blood routine parameters and (E) serum biochemical parameters of mice in the control group and H-HBXN2020 group. Data were shown as mean values ± SEM (n = 3). (F) H&E stained colon sections in the different groups. Scale bar: 200 μm.

Revised version: (A) Body weights changes of mice during gavage with B. velezensis HBXN2020 spores. Mice were treated with sterile PBS (control group) or low-dose (L-HBXN2020 group), medium dose (M-HBXN2020 group), and high-dose (H-HBXN2020 group) of B. velezensis HBXN2020 spores. Weighing and gavage were performed once every two days during the experimental period (15 days). Data were shown as mean values ± SEM (n = 5). (B) The mRNA levels of inflammatory cytokines in the colon of mice measured by RT-qPCR. Data were shown as mean values ± SEM (n = 5). (C) The mRNA levels of barrier protein ZO-1, occludin, claudin and Muc2 in the colon of mice measured by RT-qPCR. Data were shown as mean values ± SEM (n = 5). (D) Major blood routine parameters and (E) serum biochemical parameters of mice in the control group and H-HBXN2020 group. Data were shown as mean values ± SEM (n = 3). (F) H&E stained colon sections in the different groups. Scale bar: 200 μm.

Earlier version: Oral administration of B. velezensis HBXN2020 spores alleviated S. Typhimurium-induced colitis

Revised version: Oral administration of B. velezensis HBXN2020 spores alleviated infection by S. Typhimurium

Earlier version: Testing both solid and liquid co-culture methods to see if B. velezensis HBXN2020 might suppress S. Typhimurium ATCC14028 (STm), the findings indicated that B. velezensis HBXN2020 decreased the development of STm in a dose-dependent manner (Fig. 5A, S5). Next, the therapeutic potential of B. velezensis HBXN2020 was evaluated in an STm-induced mouse colitis model. At days 1, 3, and 5 after STm infection, mice in the STm + HBXN2020 group were orally administered B. velezensis HBXN2020 spores by gavage (Fig. 5B), while the control group and STm + PBS group were treated with sterile PBS. Also, the number of excreted STm in feces was counted daily at the designated time points. As shown in Fig. 5C, there was a significant and continuous reduction in the number of STm in feces following treatment with B. velezensis HBXN2020 spores, while the number of B. velezensis HBXN2020 viable bacteria in feces is also gradually decreasing (Fig. S6A). The number of STm in mouse feces was reduced by 10 to 50 times compared to the STm + PBS group from day 2 to 7 post-treatment. B. velezensis HBXN2020 spore treatment also resulted in a 5 to 10-fold reduction in the number of STm in the ileum, cecum, and colon, respectively (Fig. 5D-F). The increased inhibition of STm after B. velezensis HBXN2020 spore treatment also resulted in less weight loss and a marked reduction (P < 0.001) in DAI scores in STm + HBXN2020 group mice (Fig. 5G-H). In addition, we also measured the colon length of the mice and found that the STm + PBS group had a significantly shorter colon length (P < 0.001) than the control group. However, compared to the STm + PBS group, the STm + HBXN2020 group exhibited a longer colon (P < 0.05) (Fig. 5I-J).

Revised version: Testing both solid and liquid co-culture methods to see if B. velezensis HBXN2020 might suppress S. Typhimurium ATCC14028 (STm), the findings indicated that B. velezensis HBXN2020 decreased the development of STm in a dose-dependent manner (Fig. 5A, S6). Next, the therapeutic potential of B. velezensis HBXN2020 was evaluated in an STm-infected mouse model. At days 1, 3, and 5 after STm infection, mice in the STm+HBXN2020 group and HBXN2020 group were orally administered B. velezensis HBXN2020 spores by gavage, STm+CIP group were given oral ciprofloxacin, while the control group and STm+PBS group were orally treated with sterile PBS (Fig. 5B). Also, the number of excreted STm in feces was counted daily at the designated time points. As shown in Fig. 5C, there was a significant and continuous reduction in the number of STm in feces following treatment with B. velezensis HBXN2020 spores or ciprofloxacin, while the number of B. velezensis HBXN2020 viable bacteria in feces is also gradually decreasing (Fig. S7A). The number of STm in the mice feces of STm+HBXN2020 group and STm+CIP group exhibited a reduction ranging from 0.12 to 1.18 logs and 0.76 to 1.81 logs respectively, compared with the STm+PBS group, over the period from day 2 to day 7 post-treatment. Moreover, treatment with B. velezensis HBXN2020 spore or ciprofloxacin, resulted in a 1.06, 1.69, 1.14 logs and 1.70, 2.15, 1.48 logs reduction in the number of STm in the cecum, colon and ileum, respectively (Fig. 5D-E, S7B). The weight loss (p < 0.05) and Disease Activity Index (DAI) scores (p < 0.01) were significantly reduced in mice after B. velezensis HBXN2020 spore treatment or ciprofloxacin treatment compared to that in the PBS treatment group (Fig. 5F-G). In addition, we also measured the colon length of the mice and found that the STm+PBS group had a significantly shorter colon length (p < 0.001) than the control group and HBXN2020 group. However, compared to the STm+PBS group, the STm+HBXN2020 group (p < 0.05) and STm+CIP group exhibited a longer colon (Fig. 5H-I).

Earlier version: Oral B. velezensis HBXN2020 spores alleviated S. Typhimurium-induced experimental mouse colitis.

Revised version: Oral B. velezensis HBXN2020 spores alleviated infection by S. Typhimurium.

Earlier version: (A) In vitro Bacterial competition between STm and B. velezensis HBXN2020. STm were co-incubated with B. velezensis HBXN2020 at various ratios at 37°C with shaking. The growth of STm was reflected by bacterial counting per hour. (B) Experimental design for treatment in this study. Orally administrated with either PBS, or B. velezensis HBXN2020 spores by gavage at days 1, 3, and 5 after STm infection (5× 107 CFU/mouse), respectively. All mice were euthanized at day 7 after STm infection. (C) Bacterial count of STm in mouse feces. Fecal samples were collected per day after STm infection and resuspended in sterile PBS (0.1 g of fecal resuspended in 1 mL of sterile PBS). One hundred microliters of each sample performed a serial of 10-fold dilutions and spread on selective agar plates (50 µg mL-1 kanamycin) and incubated at 37°C for 12 h before bacterial counting. The bacterial loads of STm in (D) ileum, (E) cecum, and (F) colon. The cecum and colon were harvested and then homogenized. Statistical signifcance was evaluated using Student’s t-test (*, P < 0.05, **, P < 0.01, and ***, P < 0.001). (G) Daily body weight changes and (H) daily disease activity index (DAI) scores of mice with different treatment groups. Data were shown as mean values ± SEM (n = 8). Statistical signifcance was evaluated using one-way analysis of variance (ANOVA) with Tukey’s multiple comparisons test. *, P < 0.05, **, P < 0.01, and ***, P < 0.001 relative to Control group; #, P < 0.05, ##, P < 0.01, ###, P < 0.001 relative to STm + HBXN2020 group. (I) Colonic tissue images. (J) The length of the colon from per group (n = 8). Statistical signifcance was evaluated using one-way analysis of variance (ANOVA) with Tukey’s multiple comparisons test (*, P < 0.05, **, P < 0.01, and ***, P < 0.001).

Revised version: (A) In vitro Bacterial competition between STm and B. velezensis HBXN2020. STm were co-incubated with B. velezensis HBXN2020 at various ratios at 37°C with shaking. The growth of STm was reflected by bacterial counting per hour. (B) Experimental design for treatment in this study. Orally administrated with either sterile PBS, B. velezensis HBXN2020 spores or ciprofloxacin by gavage at days 1, 3, and 5 after STm (5× 107 CFU/mouse) infection, respectively. All mice were euthanized at day 7 after STm infection. (C) Bacterial count of STm in mouse feces. Fecal samples were collected per day after STm infection and resuspended in sterile PBS (0.1 g of fecal resuspended in 1 mL of sterile PBS). One hundred microliters of each sample performed a serial of 10-fold dilutions and spread on selective agar plates (50 µg mL-1 kanamycin) and incubated at 37°C for 12 h before bacterial counting. The bacterial loads of STm in (D) cecum and (E) colon. The cecum and colon were harvested and then homogenized. Statistical signifcance was evaluated using Student’s t-test (*, p < 0.05, **, p < 0.01, and ***, p < 0.001). (F) Daily body weight changes and (G) daily disease activity index (DAI) scores of mice with different treatment groups. Data were shown as mean values ± SEM (n = 6). Statistical signifcance was evaluated using one-way analysis of variance (ANOVA) with Tukey’s multiple comparisons test (*, p < 0.05, **, p < 0.01, and ***, p < 0.001). (H) Colonic tissue images. (I) The length of the colon from per group (n = 6). Statistical signifcance was evaluated using one-way analysis of variance (ANOVA) with Tukey’s multiple comparisons test (*, p < 0.05, **, p < 0.01, and ***, p < 0.001).

Added in the revised version: Histological analysis further showed that in the control group and HBXN2020 group, the colon epithelial cells and crypt structure were intact with neat villi, while STm+PBS group showed significant histological damage, including erosion or loss of the intestinal epithelium, crypt destruction, and inflammatory cell infiltration in the colonic tissues (Fig. 6A). Compared with the STm+PBS group, B. velezensis HBXN2020 and ciprofloxacin protected the mucosal architecture and the loss of intestinal epithelial cells, and reduced inflammatory cell infiltration (Fig. 6A).

Removed in the revised version: To further assess the impact of B. velezensis HBXN2020 on the intestinal tract, we histologically analyzed their colonic tissues. The results showed that mice treated with B. velezensis HBXN2020 spores exhibited a lower histology score (total scoring) (P < 0.05) (Fig. 6A, S6B) than mice not treated with the spores and were similar to that of healthy mice (control group). The colon tissue of mice not treated with B. velezensis HBXN2020 spores displayed higher histological scores, with significant crypt deformation, severe mucosal damage, and obvious inflammatory cell infiltration. These findings suggested that B. velezensis HBXN2020 can alleviate colon tissue damage. Moreover, we measured the mRNA levels of inflammatory cytokines, including tumor necrosis factor-α (TNF-α), interleukin-1 β (IL-1 β), interleukin-6 (IL-6), and interleukin-10 (IL-10) in colon tissue using RT-qPCR analysis. As shown in Fig. 6B-D, the mRNA levels of TNF-α (P < 0.01), IL-1β (P < 0.05), and IL-6 (P < 0.05) in the colon tissue of STm-infected mice were significantly upregulated, indicating the successful establishment of the STm-induced colitis model. As expected, treatment with B. velezensis HBXN2020 spores led to a significant reduction in the mRNA levels of these pro-inflammatory cytokines, while the mRNA levels of the anti-inflammatory cytokine IL-10 were significantly increased (P < 0.01) (Fig. 6E), suggesting that B. velezensis HBXN2020 spores can alleviate the inflammatory reaction of the colon. Moreover, the mRNA expression levels of colonic epithelial barrier proteins, including ZO-1, ocludin, claudin, and Muc2, were also measured by RT-qPCR. Compared with the PBS-treated group, the B. velezensis HBXN2020 spores treatment group exhibited significantly upregulated mRNA expression levels of ZO-1 (P < 0.05), occludin (P < 0.05), claudin (P < 0.05), and Muc2 (P < 0.01) in the colon tissue (Fig. 6F-I), further demonstrating the ability of B. velezensis HBXN2020 to reduce damage to the colon epithelial barrier.

Oral B. velezensis HBXN2020 spores attenuated colonic damage and inflammatory reaction.

Earlier version: (A) H&E stained colon tissue sections. Scale bar: 200 μm. The relative gene expression levels of TNF-α (B), IL-1β (C), IL-6 (D), and IL-10 (E) were detected by RT-qPCR. Data were shown as mean values ± SEM (n = 6). The relative gene expression levels of ZO-1 (F), ocludin (G), claudin (H), and Muc2 (I) in colon tissue were detected by RT-qPCR. Data were shown as mean values ± SEM (n = 6). Statistical signifcance was evaluated using one-way analysis of variance (ANOVA) with Tukey’s multiple comparisons test (*, P < 0.05, **, P < 0.01, and ***, P < 0.001).

Revised version: (A) H&E stained colon tissue sections. Scale bar: 50 μm. The mRNA levels of (B) TNF-α, (C) IL-1β, (D) IL-6, and (E) IL-10 were detected by RT-qPCR. Data were shown as mean values ± SEM (n = 5). The mRNA levels of (F) ZO-1, (G) occludin, (H) claudin, and (I) Muc2 in colon tissue were detected by RT-qPCR. Data were shown as mean values ± SEM (n = 5). Statistical signifcance was evaluated using one-way analysis of variance (ANOVA) with Tukey’s multiple comparisons test (*, p < 0.05, **, p < 0.01, and ***, p < 0.001).

Added in the revised version: To further assess the impact of B. velezensis HBXN2020 on intestinal infammatory response, the mRNA levels of inflammatory cytokines in the colonic tissues were measured. As shown in Fig. 6B-D, the mRNA levels of TNF-α (p < 0.05), IL-1β (p < 0.001), IL-6 (p <0.001), and IL-10 (p <0.001) in the colon tissue were significantly increased in the STm+PBS group compared to the control group. B. velezensis HBXN2020 spore treatment signifcantly decreased the mRNA levels of TNF-α (p < 0.05), IL-1β (p < 0.001), IL-6 (p < 0.01), and increased the IL-10 level, compared with the STm+PBS group. In contrast, ciprofloxacin treatment failed to signifcantly decrease the mRNA levels of IL-6 in the colon of STm-infected mice (Fig. 6D). Similarly, ciprofloxacin treatment drastically dampened the mRNA levels of IL-10 in the colon of STm-infected mice (Fig. 6E).

Added in the revised version: To assess the infuence of B. velezensis HBXN2020 on the intestinal barrier integrity, the mRNA levels of tight junction proteins in the colonic tissues were measured. As presented in Fig. 6, compared with control group, the mRNA levels of tight junction protein in the STm+PBS group was significantly reduced, including ZO-1 (p < 0.05), occludin (p < 0.01), claudin, and Muc2 (p < 0.05) (Fig. 6F-I). Moreover, the levels of tight junction protein transcription were remarkably elevated in the STm+HBXN2020, STm+CIP, and HBXN2020 group, including ZO-1 (p < 0.001, p < 0.05 and p < 0.05, respectively), occludin (p < 0.05, p < 0.05 and p < 0.01, respectively), and claudin (p < 0.001, p = 0.0857 and p < 0.01, respectively), except for claudin (p = 0.0857) in group STm+CIP, compared with that of the STm+PBS group (Fig. 6F-I). In addition, we also detected the transcription level of Muc2 in colon tissue, which mainly exists in and around goblet cells and is an important component of the mucus layer. The transcription levels of Muc2 were dramatically increased in the STm+HBXN2020 group compared with the STm+PBS group (Fig. 6I). Conversely, the mRNA levels of Muc2 were significantly decreased in the HBXN2020 group and STm+CIP group (Fig. 6I).

Added in the revised version: Next, we further explore the impact of B. velezensis HBXN2020 on the intestinal microbiota composition of STm-treated mice by 16S rRNA gene high throughout sequencing. As shown in Fig. 7, alpha-diversity analysis revealed that both the richness and diversity (calculated in Sobs, Chao, Shannon, and Simpson indexes) were lower in STm+PBS group than control group, HBXN2020 group, STm+HBXN2020 group, and STm+CIP group (Fig. 7A-D). Although an increasing trend of alpha diversity was observed in STm+CIP group, but without statistical difference was achieved as compared with other groups. Principal Components analysis (PCA) based on Bray-Curtis distance showed that a separation in the gut microbiota structure among control and STm+PBS group. When B. velezensis HBXN2020 spores were supplemented, the gut microbiota composition of the STm+HBXN2020 group, HBXN2020 group and the control group were closer together (R2= 0.1616, P = 0.006, Fig. 7E).

Removed in the revised version: Next, we used 16S rRNA gene sequencing to investigate the impact of HBXN2020 spores on the makeup of the intestinal microbiota in mice treated with STm. The alpha diversity of colonic microbiota was evaluated using Sobs, Chao1, Shannon, and Simpson indices (OTU levels). Compared with the control and HBXN2020 group, the Sobs, Chao1, and Shannon indices of the STm + PBS group clearly decreased (Fig. 7A-C), while the Simpson index significantly increased (P < 0.05, Fig. 7D). Principle component analysis (PCA) based on Bray-Curtis distance showed that compared with the STm + PBS group, the gut microbiota composition of the STm + HBXN2020 group and the control group were more closer together (R= 0.3988, P = 0.001, Fig. 7E). We analyzed the community composition of colonic microbiota at the phylum and genus level, and the results revealed that the composition of the gut microbiota changed markedly after STm-induced mice. At the phylum level, the STm + PBS group increased the abundance of Verrucomicrobiota and Proteobacteria and decreased the abundance of Firmicutes, Patescibacteria, Desulfobacterota, and Actinobacteriota compared to the Control group (Fig. 7F). The HBXN2020 group increased the abundance of Bacteroidota and Firmicutes and decreased the abundance of Verrucomicrobiota compared with the Control group. On the contrary, the relative abundance of Firmicutes, Patescibacteria, Desulfobacterota, and Actinobacteriota in the STm + HBXN2020 group was significantly higher than that in the STm + PBS group. At the genus level, infection with S. Typhimurium ATCC14028 in the STm + PBS group dramatically reduced the relative abundance of norank_f_Muribaculaceae, Desulfovibrio and Lactobacillus (Fig. 7G, H-I) and enhanced the abundance of Bacteroides, Escherichia-Shigella, Salmonella, Alistipes, and Enterococcus compared to the Control group (Fig. 7G, J-K). The HBXN2020 group increased the abundance of norank_f_Muribaculaceae and Alloprevotella compared to the Control group. In contrast, the STm + HBXN2020 group significantly improved the relative abundance of norank_f_Muribaculaceae and Lactobacillus compared with the STm + PBS group (P < 0.01, P < 0.001, Fig. 7H-I). In addition, LEfSe analysis was also performed to identify specialized microbial communities between the STm + HBXN2020 group and the STm + PBS group and found that there was a clear difference in gut microbiota composition between them. The microbiota that played a major role in the STm + PBS group were harmful Enterobacterales (Fig. 7L), while the STm + HBXN2020 group was Candidatus_Saccharimonas.

Oral B. velezensis HBXN2020 spores regulated the composition of intestinal microbiota.

Earlier version: The α-diversity of the gut microbiota, determined by the (A) Sobs, (B) Chao1, (C) Shannon, and (D) Simpson diversity index. Data were shown as mean values ± SEM (n = 5). (E) The PCA plot showed the β-diversity of the gut microbiota based on Bray-Curtis distance at the OTU level. Heatmap of the community composition of colonic microbiota at the phylum (top 10 phyla) (F) and genus (top 30 genera) (G) level. Relative abundance of selected taxa (H) norank_f_Muribaculaceae, (I) Lactobacillus, (J) Bacteroides, and (K) Escherichia-Shigella. (L) Analysis of differences in the microbial communities by LEfSe (linear discriminant analysis (LDA score > 3.5) among different groups. Signifcance was evaluated by the Kruskal-Wallis test or ANOVA with Tukey’s multiple comparisons test (*, P < 0.05, **, P < 0.01, and ***, P < 0.001).

Revised version: The α-diversity of the gut microbiota, determined by the (A) Sobs, (B) Chao, (C) Shannon, and (D) Simpson index. Data were shown as mean values ± SEM (n = 5). (E) The PCA plot showed the β-diversity of the gut microbiota based on Bray-Curtis distance at the OTU level. Heatmap of the community composition of colonic microbiota at the phylum (top 15 phyla) (F) and genus (top 30 genera) (G) level. Relative abundance of selected taxa (H) norank_f_Muribaculaceae, (I) Lactobacillus, (J) Akkermansia, and (K) Enterococcus. (L) Analysis of differences in the microbial communities by LEfSe (linear discriminant analysis, LDA score > 2) among different groups. Signifcance was evaluated by one-way analysis of variance (ANOVA) with Tukey’s multiple comparisons test (*, p < 0.05, **, p < 0.01, and ***, p < 0.001).

Added in the revised version: We analyzed the community composition of colonic microbiota at the phylum and genus level, and the results revealed that the composition of the gut microbiota changed markedly after STm-infected mice. At the phylum level, Bacteroidetes, Firmicutes, Verrucomicrobiota and Proteobacteria were predominant phyla in the fecal microbiota (Fig. 7F, S8A). At the genus level, infection with STm in the STm+PBS group dramatically reduced the relative abundance of norank_f_Muribaculaceae, Lactobacillus and Akkermansia (Fig. 7G, H-J) and enhanced the abundance of Enterococcus, Salmonella, Bacteroides, Alloprevotella and Escherichia-Shigella compared to the control group (Fig. 7G, K, S8B). In contrast, the STm+HBXN2020 group and HBXN2020 group significantly improved the relative abundance of norank_f_Muribaculaceae (p = 0.49, p < 0.05,) and Lactobacillus (p < 0.05, p < 0.001) compared with the STm+PBS group (Fig. 7H-I). However, ciprofloxacin treatment significantly decreased the relative abundance of Lactobacillus (p < 0.05) and Akkermansia (p < 0.05) compared with the STm+HBXN2020 group (Fig. 7I-J). Differentially abundant fecal bacterial taxa in STm-treated mice in response to B. velezensis HBXN2020 and ciprofloxacin were identified by LEfSe analysis (Fig. 7L). Additionally, we found that two bacterial genera including Enterococcus and Streptococcus were enriched in the STm+PBS group, while the other three taxa were enriched including Akkermansia, Family_XIII_AD3011_group, and Corynebacterium in the STm+HBXN2020 group, Parabacteroides, NK4A214_group, Turicibacter, unclassified_o Oscillospirales, and norank_f Eubacterium_coprostanoligenes_group in the STm+CIP group, and norank_f_Muribaculaceae, Rikenellaceae_RC9_gut_group and Dubosiella in the HBXN2020 group (Fig. 7L).

Earlier version: Prophylactic B. velezensis HBXN2020 spores alleviated S. Typhimurium-induced colitis

Revised version: Prophylactic B. velezensis HBXN2020 spores alleviated infection by S. Typhimurium

Earlier version: Based on the improved therapeutic efficacy of B. velezensis HBXN2020 in the treatment of STm-induced colitis in mice, we further explored its potential for disease prevention by evaluating pretreatment. Mice in the PBS + STm group and HBXN2020 + STm group were given the same amount of sterile PBS and B. velezensis HBXN2020 spores by gavage one week in advance. As shown in Fig. 8, compared with the PBS + STm group, oral administration of B. velezensis HBXN2020 spores as a preventive measure remarkably alleviated STm-induced colitis, including weight loss of mice (P < 0.01) (Fig. 8B), a significant reduction in DAI (P < 0.001) (the comprehensive score of weight loss, stool consistency, and blood in the feces, Fig. 8C), and the prevention of colon length shortening (Fig. 8D-E). However, the number of B. velezensis HBXN2020 viable bacteria in feces is also gradually decreasing with time prolonging (Fig. S8A). Furthermore, compared with the PBS + STm group, oral administration of B. velezensis HBXN2020 spores not only reduced the number of STm in mouse feces (Fig. 8F) but also decreased STm colonization in the ileum, cecum, and colon (Fig. 8G-H, S8B).

Revised version: Based on the improved therapeutic efficacy of B. velezensis HBXN2020 in the treatment of STm-infected mice, we further explored its potential for disease prevention by evaluating pretreatment. Mice in the PBS+STm group, HBXN2020+STm group, and CIP+STm group were given the same amount of sterile PBS, B. velezensis HBXN2020 spores or ciprofloxacin by gavage one week in advance. As shown in Fig. 8, compared with the PBS+STm group, oral administration of B. velezensis HBXN2020 spores or ciprofloxacin as a preventive measure remarkably alleviated infection by STm, including weight loss of mice (p < 0.01) (Fig. 8B), a significant reduction in DAI (p < 0.001) (the comprehensive score of weight loss, stool consistency, and blood in the feces, Fig. 8C), and the prevention of colon length shortening (Fig. 8D-E). Meanwhile, the number of B. velezensis HBXN2020 viable bacteria in feces is also gradually decreasing with time prolonging (Fig. S9A). Furthermore, compared with the PBS+STm group, oral administration of B. velezensis HBXN2020 spores or ciprofloxacin not only reduced the number of STm in mouse feces (Fig. 8F) but also decreased STm colonization in the ileum, cecum, and colon (Fig. 8G-H, S9B).

Earlier version: Prophylactic B. velezensis HBXN2020 spores attenuated the symptoms of S. Typhimurium-induced mouse colitis.

Revised version: Prophylactic B. velezensis HBXN2020 spores attenuated the symptoms of S. Typhimurium-infected mouse.

Earlier version: (A) Experimental design for treatment in this study. At days 1, 3, 5, and 7, each mouse in the HBXN2020 + STm group and PBS + STm group were received 200 μL (1×108 CFU/mouse) of B. velezensis HBXN2020 spores or sterile PBS by gavage. Then, mice in PBS + STm group and HBXN2020 + STm group were orally inoculated with 200 μL (5×107 CFU/mouse) of STm on day 7. On day 12, all mice were euthanized. (B) Daily body weight changes and (C) daily disease activity index (DAI) scores of mice with different groups following STm treatment. Data were shown as mean values ± SEM (n = 8). Statistical signifcance was evaluated using one-way analysis of variance (ANOVA) with Tukey’s multiple comparisons test. *, P < 0.05, **, P < 0.01, and ***, P < 0.001 relative to Control group; #, P < 0.05, ##, P < 0.01, ###, P < 0.001 relative to HBXN2020 + STm group. (D) Colonic tissue images. (E) The length of the colon from per group (n = 8). (F) Bacterial count of STm in mouse feces. Fecal samples were collected every day after STm infection and resuspended in sterile PBS (0.1 g of fecal resuspended in 1 mL of sterile PBS). One hundred microliters of each sample performed a serial of 10-fold dilutions and spread on selective agar plates (50 µg mL-1 kanamycin) and incubated at 37°C for 12 h before bacterial counting. The bacterial loads of STm in (G) cecum and (H) colon. The cecum and colon were harvested and then homogenized. Statistical signifcance was evaluated using Student’s t-test (*, P < 0.05, **, P < 0.01, and ***, P < 0.001).

Revised version: (A) Experimental design for treatment in this study. At days 1, 3, 5, and 7, each mouse in the HBXN2020+STm group, CIP+STm group, and PBS+STm group were received 200 μL of B. velezensis HBXN2020 spores (1×108 CFU/mouse), ciprofloxacin or sterile PBS by gavage, respectively. Then, mice in PBS+STm group, HBXN2020+STm group and CIP+STm group were orally inoculated with 200 μL (5×107 CFU/mouse) of STm on day 7, respectively. On day 12, all mice were euthanized. (B) Daily body weight changes and (C) daily disease activity index (DAI) scores of mice with different groups following STm treatment. Data were shown as mean values ± SEM (n = 6). Statistical signifcance was evaluated using one-way analysis of variance (ANOVA) with Tukey’s multiple comparisons test (*, p < 0.05, **, p < 0.01, and ***, p < 0.001). (D) Colonic tissue images. (E) The length of the colon from per group (n = 6). (F) Bacterial count of STm in mouse feces. Fecal samples were collected every day after STm infection and resuspended in sterile PBS (0.1 g of fecal resuspended in 1 mL of sterile PBS). One hundred microliters of each sample performed a serial of 10-fold dilutions and spread on selective agar plates (50 µg mL-1 kanamycin) and incubated at 37°C for 12[h before bacterial counting. The bacterial loads of STm in (G) cecum and (H) colon. The cecum and colon were harvested and then homogenized. Statistical signifcance was evaluated using Student’s t-test (*, p < 0.05, **, p < 0.01, and ***, p < 0.001).

Earlier version: Histological analysis further revealed that oral administration of B. velezensis HBXN2020 significantly reduced the infiltration of inflammatory cells, mucosal damage and the overall histological score of colon tissue (Fig. 9A, S8C). Additionally, the mRNA expression levels of proinflammatory cytokines, including TNF-α (P < 0.05), IL-1β (P < 0.05) and IL-6 (P < 0.01), in the colon tissue of mice pretreated with B. velezensis HBXN2020 spores were significantly decreased (Fig. 9B-D), while the mRNA levels of anti-inflammatory cytokine (IL-10) were significantly higher (P < 0.05) (Fig. 9E) than mice pretreated with PBS. In addition, compared to the group pretreated with PBS, the group pretreated with B. velezensis HBXN2020 spores showed a significant increase in mRNA expression levels of ZO-1 (P < 0.001), occludin (P < 0.05) and claudin (P < 0.05) in the colon tissue of mice (Fig. 9F-H). While there was no significant difference in Muc2 expression (p=0.252) (Fig. 9I), the results were consistent with those of the therapeutic experiments. Overall, these findings suggest that B. velezensis HBXN2020 can effectively alleviate the symptoms of STm-induced colitis and the associated damage to colon tissue.

Revised version: Histological analysis further revealed that prophylactic B. velezensis HBXN2020 spore suppressed STm induced loss of the intestinal epithelium, inflammatory cell infiltration and colonic mucosa damage (Fig. 9A), while prophylactic ciprofloxacin failed to inhibit the loss of colonic epithelial cells. We next measured the mRNA levels of inflammatory cytokines and tight junction proteins in colon tissue in the prophylactic animal experiment. As with the treatment test, prophylactic B. velezensis HBXN2020 spores largely attenuated the mRNA levels of TNF-α (p < 0.001), IL-1β (p < 0.001) and IL-6 (p < 0.001), and significantly increased levels of IL-10 (p < 0.05), in the colon tissue of mice compared with the PBS+STm group (Fig. 9B-E). Similarly, prophylactic ciprofloxacin significantly reduced the mRNA levels of TNF-α (p < 0.01) and IL-1β (p < 0.001), and drastically elevated levels of IL-10 (p < 0.001) in the colon tissue of mice compared with the PBS+STm group (Fig. 9B-C, E). Also, prophylactic B. velezensis HBXN2020 spores signifcantly increased the transcription levels of ZO-1 (p < 0.001), occludin (p < 0.001), claudin (p < 0.05), and Muc2 (p < 0.001) compared with that of the PBS+STm group (Fig. 9F-I). Notably, the mRNA levels of ZO-1 (p < 0.001), occludin (p < 0.05), and Muc2 (p < 0.05) in the CIP+STm group were significantly lower than those in the HBXN2020+STm group. Overall, these results suggested that prophylactic B. velezensis HBXN2020 spores was capable of alleviating STm-induced colonic injury, and inflammation.

Prophylactic B. velezensis HBXN2020 spores attenuated colonic damage and inflammatory reaction.

Earlier version: (A) H&E stained colon tissue sections. Scale bar: 200 μm. The relative gene expression levels of TNF-α (B), IL-1β (C), IL-6 (D), and IL-10 (E) were detected by RT-qPCR. Data were shown as mean values ± SEM (n = 6). The relative gene expression levels of ZO-1 (F), ocludin (G), claudin (H), and Muc2 (I) in colon tissue were detected by RT-qPCR. Data were shown as mean values ± SEM (n = 6). Statistical signifcance was evaluated using one-way analysis of variance (ANOVA) with Tukey’s multiple comparisons test (*, P < 0.05, **, P < 0.01, and ***, P < 0.001).

Revised version: (A) H&E stained colon tissue sections. Scale bar: 50 μm. The mRNA levels of (B) TNF-α, (C) IL-1β, (D) IL-6, and (E) IL-10 were detected by RT-qPCR. Data were shown as mean values ± SEM (n = 5). The mRNA levels of (F) ZO-1, (G) occludin, (H) claudin, and (I) Muc2 in colon tissue were detected by RT-qPCR. Data were shown as mean values ± SEM (n = 6). Statistical signifcance was evaluated using one-way analysis of variance (ANOVA) with Tukey’s multiple comparisons test (*, p < 0.05, **, p < 0.01, and ***, p < 0.001).

Earlier version: Next, we examined the impact of prophylactic B. velezensis HBXN2020 on the intestinal microbiota composition of STm-treated mice. Sobs, Chao1 index and Shannon index have changed to some extent between groups, but differences were not significant (Fig. 10A-C). Compared with the Control group, the Simpson index significantly decreased in the PBS + STm group (P < 0.05, Fig. 10D). However, compared to the PBS + STm group, Simpson index in the HBXN2020 + STm group did not change (Fig. 10D). PCA analysis showed significant separation between control and PBS + STm groups, while the HBXN2020 + STm group, HBXN2020 group, and the control group were clustered together (Fig. 10E). Next, we analyzed microbial community composition at the phylum and genus levels. At the phylum level, the PBS + STm group increased the abundance of Bacteroidota, Proteobacteria, Actinobacteriota, and Deferribacterota and reduced the abundance of Verrucomicrobiota, Firmicutes, Patescibacteria, and Desulfobacterota compared with the Control group (Fig. 10F). The HBXN2020 group increased the abundance of Bacteroidota compared to the control group. In contrast, compared to the PBS + STm group, the abundance of Firmicutes in the HBXN2020+STm group was significantly increased (P < 0.05). At the genus level, a depletion of Lactobacillus and Akkermansia (Fig. 10G, H-I) and an enrichment of Alistipes, Bacteroides, Escherichia-Shigella, Alloprevotella, Helicobacter, Enterococcus, and Salmonella were observed in PBS + STm group compared with the control (Fig. 10G, J-M). The enrichment with norank_f_Muribaculaceae and Alloprevotella was observed in the HBXN2020 group compared to the control. In contrast, an enrichment of Lactobacillus, Akkermansia, and Lachnospiraceae_NK4A136_group and a sharp decrease of Alistipes, Bacteroides, Escherichia-Shigella, Salmonella, and Enterococcus were observed in HBXN2020 + STm group compared to the PBS + STm group. LEfSe analysis showed that the abundance of Lactobacillus in the intestinal microbiota of mice was significantly increased after B. velezensis HBXN2020 treatment, while the abundance of harmful Enterobacterales in the untreated group was clearly increased (Fig. 10N).

Revised version: Next, we examined the impact of prophylactic B. velezensis HBXN2020 on the intestinal microbiota composition of STm-treated mice. Sobs, Chao and Shannon index were signifcantly increased in the HBXN2020, PBS+STm group, HBXN2020+STm group and CIP+STm group compared with the control group (Fig. 10A-C), while the Simpson index (not statistically different) significantly decreased (Fig. 10D). PCA analysis showed significant separation between control and PBS+STm groups (R2=0.3186, p=0.014, Fig. 10E). When B. velezensis HBXN2020 spores or ciprofloxacin were supplemented, the community clustering was significantly similar to that of the control group rather than the PBS+STm group (Fig. 10E). Meanwhile, the overall fecal bacterial composition of mice at the phylum and genus level in all groups was similar to that of the treatment experiment. The fecal microbiota was dominated by phyla Firmicutes, Bacteroidetes, Verrucomicrobiota, Patescibacteria, and Actinobacteria in all five groups (Fig. 10F, S10A). At the genus level, PBS+STm group exhibited lower relative abundances of Lactobacillus and Akkermansia (Fig. 10G, H-I), and higher relative abundances of Alistipes, Bacteroides, Escherichia-Shigella, Enterococcus, and Alloprevotella, compared with the control (Fig. 10G, J-M, S10B). Moreover, compared with the PBS+STm group, following supplementation with B. velezensis HBXN2020 spores or ciprofloxacin, Lactobacillus significantly increased, and the gut microbiota was restored to a composition similar to the control group. It was noteworthy that six bacterial genera including Alistipes, Odoribacter, and norank_f Oscillospiraceae were enriched in PBS+STm group, while other three taxa (e.g., Lactobacillus and Akkermansia) were enriched in the control group (Fig. 10N). Moreover, we found that Candidatus_Arthromitus was enriched in the HBXN2020+STm group, Dubosiella was enriched in the CIP+STm group, and Faecalibaculum, Rikenellaceae_RC9_gut_group, and unclassified_f Lactobacillaceae were enriched in the HBXN2020 group (Fig. 10N). Above all, these results indicated that prophylactic B. velezensis HBXN2020 spores regulated the gut microbiota composition, leading to the potential to attenuate the STm-induced dysbiosis.

Prophylactic B. velezensis HBXN2020 spores regulated the composition of gut microbiota.

Earlier version: (A to D) Alpha diversity of the intestinal microbiota. Data were shown as mean values ± SEM (n = 5). (E) The PCA plot showed the β-diversity among different microbial community groups based on Bray-Curtis distance at the OTU level. Heatmap of the community composition of colonic microbiota at the phylum (top 10 phyla) (F) and genus (top 30 genera) (G) level. Relative abundance of selected taxa (H) Lactobacillus, (I) Akkermansia, (J) Alistipes, (K) Bacteroides, (L) Escherichia-Shigella, and (M) Enterococcus. Data were shown as mean values ± SEM (n = 5). (N) Analysis of differences in the microbial taxa by LEfSe (LDA score > 3.5) in different groups. Signifcance was evaluated by the Kruskal-Wallis test or ANOVA with Tukey’s multiple comparisons test (*, P < 0.05, **, P < 0.01, and ***, P < 0.001).

Revised version: (A to D) Alpha diversity of the intestinal microbiota. Data were shown as mean values ± SEM (n = 5). (E) The PCA plot showed the β-diversity among different microbial community groups based on Bray-Curtis distance at the OTU level. Heatmap of the community composition of colonic microbiota at the phylum (top 15 phyla) (F) and genus (top 30 genera) (G) level. Relative abundance of selected taxa (H) Lactobacillus, (I) Akkermansia, (J) Alistipes, (K) Bacteroides, (L) Escherichia-Shigella, and (M) Enterococcus. Data were shown as mean values ± SEM (n = 5). (N) Analysis of differences in the microbial taxa by LEfSe (LDA score > 2) in different groups. Signifcance was evaluated by one-way analysis of variance (ANOVA) with Tukey’s multiple comparisons test (*, p < 0.05, **, p < 0.01, and ***, p < 0.001).

Discussion

S. Typhimurium is an important intestinal pathogen that can cause invasive intestinal diseases such as bacterial colitis (Herp et al., 2019; Schultz et al., 2017). Antibiotics are frequently used to treat colitis caused by bacteria; nevertheless, in recent years, their abuse has greatly increased bacterial resistance, leading to serious environmental pollution. A recent study showed that antibiotic-resistance genes in probiotics could be transmitted to the intestinal microbiota, which may threaten human health (Crits-Christoph et al., 2022). It is crucial to consider the source of probiotics. Strong environmental adaptability and disease resistance make black pigs an outstanding household breed in China (Yang et al., 2022). In this study, B. velezensis HBXN2020 was isolated from the free-range feces of black piglets in a mountain village in Xianning City (Hubei, China) and exhibited excellent antibacterial activity. Otherwise, in vitro tolerance assays showed that B. velezensis HBXN2020 spores have good tolerance to high temperature, strong acids, bile salts, as well as simulated gastric and intestinal fluid, which was similar to previous research results (Du et al., 2022). Furthermore, studies for antibiotic susceptibility revealed that B. velezensis HBXN2020 is not resistant to antibiotics.

Probiotics are distinct from food or medications in that they can cause toxins or infections in the body when ingested, and they are living when swallowed (Sanders et al., 2010; Snydman, 2008). Therefore, assessing the biosafety of probiotics is crucial. In this study, the safety of B. velezensis HBXN2020 was evaluated by measuring the body weight, intestinal barrier proteins, and inflammatory cytokines of the experimented mice, as well as conducting routine blood and biochemical tests. The results showed that the expression levels of ZO-1 and occludin in the colon of mice increased after treatment with different doses of B. velezensis HBXN2020 spores. It has been reported that ZO-1, occludin and claudin are the three most important tight junction proteins in intercellular connections, which play a crucial role in maintaining the intestinal epithelial barrier (Bazzoni et al., 2000; Zihni et al., 2016). Thus, it is possible to improve intestinal barrier function by upregulating ZO-1 and occludin expression levels.

Earlier version: Important aminotransferases in animals, aspartate aminotransferase (AST) and alanine aminotransferase (ALT), are regarded as critical metrics for assessing liver function damage (Gou et al., 2022; Ozer et al., 2008). Under normal circumstances, the levels of ALT and AST are in dynamic balance without notable changes. However, when the liver is damaged or becomes dysfunctional, the levels of aminotransferases (ALT and AST) increase significantly (Tang et al., 2012). Previous studies have shown that CCL4-induced liver injury strongly increases plasma ALT/AST levels (Singhal et al., 2018). Zhang et al. reported that the addition of Bacillus subtilis to chicken diets significantly decreased the serum levels of ALT and AST (Zhang et al., 2017). In this present study, the levels of ALT and AST in the probiotic-treated group were slightly lower than those in the control group, indicating that B. velezensis HBXN2020 had no negative effect on liver health in mice. The B. velezensis HBXN2020 treatment group’s blood indices were also found to be 18 comparable to those of the control group, falling within the normal reference range, according to regular blood testing.

Revised version: Important aminotransferases in animals, aspartate aminotransferase (AST) and alanine aminotransferase (ALT), are regarded as critical metrics for assessing liver function damage (Gou et al., 2022; Ozer et al., 2008). Under normal circumstances, the levels of ALT and AST are in dynamic balance without notable changes. However, when the liver is damaged or becomes dysfunctional, the levels of aminotransferases (ALT and AST) increase significantly (Tang et al., 2012). Previous studies have shown that CCL4-induced liver injury strongly increases plasma ALT/AST levels (Singhal et al., 2018). Zhang et al. reported that the addition of Bacillus subtilis to chicken diets significantly decreased the serum levels of ALT and AST (Zhang et al., 2017). In this present study, the levels of ALT and AST in the probiotic-treated group were slightly lower than those in the control group, indicating that B. velezensis HBXN2020 had no negative effect on liver health in mice. The B. velezensis HBXN2020 treatment group’s blood indices were also found to be comparable to those of the control group, falling within the normal reference range, according to regular blood testing.

Earlier version: Mice were fed with B. velezensis HBXN2020 spores orally after being challenged with S. Typhimurium ATCC14028 in order to assess the spores’ potential to alleviate colitis caused by S. Typhimurium. Our experimental results demonstrated that oral treatment with B. velezensis HBXN2020 spores could alleviate S. Typhimurium-induced colitis, as evidenced by decreased weight loss, DAI, and histological damage, which is similar to prior studies (Cao et al., 2019; Wu et al., 2022). Previous research showed that macrophages are the first line of host defense against bacterial infection. They release proinflammatory cytokines, which are critical in initiating adaptive immune responses (Cheng et al., 2014; Dinarello, 2000). Nevertheless, proinflammatory cytokines have immunological properties that can be beneficial for the host to resist the invasion of bacteria and other microbes in the surrounding environment (Liu et al., 2021). For instance, pretreatment with recombinant murine TNF-α was shown to protect mice against lethal bacterial (E. coli) infection (Cross et al., 1989). PJ-34 exerted protective effects on intestinal epithelial cells against invasive Salmonella infection by upregulating IL-6 expression through the ERK and NF-κB signal pathways (Huang, 2009).

Revised version: Mice were fed with B. velezensis HBXN2020 spores orally after being challenged with S. Typhimurium ATCC14028 in order to assess the spores’ potential to alleviate infection caused by S. Typhimurium. Our experimental results demonstrated that oral treatment with B. velezensis HBXN2020 spores could alleviate infection by S. Typhimurium, as evidenced by decreased weight loss, DAI, and histological damage, which is similar to prior studies (Cao et al., 2019; Wu et al., 2022). Previous research showed that macrophages are the first line of host defense against bacterial infection. They release proinflammatory cytokines, which are critical in initiating adaptive immune responses (Cheng et al., 2014; Dinarello, 2000). Nevertheless, proinflammatory cytokines have immunological properties that can be beneficial for the host to resist the invasion of bacteria and other microbes in the surrounding environment (Liu et al., 2021). For instance, pretreatment with recombinant murine TNF-α was shown to protect mice against lethal bacterial (E. coli) infection (Cross et al., 1989). PJ-34 exerted protective effects on intestinal epithelial cells against invasive Salmonella infection by upregulating IL-6 expression through the ERK and NF-κB signal pathways (Huang, 2009).

On the other hand, some research has demonstrated that over-activation of immune cells can result in tissue damage, organ failure, systemic or persistent inflammation, and autoimmune diseases (Karki et al., 2021; Kotas and Medzhitov, 2015). IL-10 is a recognized anti-inflammatory mediator that plays a crucial role in maintaining intestinal microbe-immune homeostasis, regulating the release of inflammatory mediators, and inhibiting proinflammatory responses of innate and adaptive immunity (Maloy and Powrie, 2011; Ouyang et al., 2011; Saraiva et al., 2020). For instance, previous studies have shown that Clostridium butyrate can induce the production of IL-10 in the intestine, thereby alleviating experimental colitis in mice (Hayashi et al., 2013). Similar results were observed in our study, where oral administration of B. velezensis HBXN2020 spores reduced the expression levels of TNF-α, IL-1β, and IL-6 while increasing the levels of IL-10 in the colon of mice. Additionally, as demonstrated by the decreased histological damage and increased expression levels of intestinal barrier proteins, oral treatment of B. velezensis HBXN2020 spores reduced the functional damage of the intestinal barrier induced by S. Typhimurium infection.

Added in the revised version: Besides host itself, environmental factors such as diet and gut microbiota have been associated with the development of STm-infected. Gut microbiota constitutes a critical bridge between environmental factors and host health, where the beneficial bacteria such as Lactobacillus in the microbiota may exert repair and anti-infammatory function (Li et al., 2022; Liu et al., 2022). Moreover, supplementation of probiotics has been shown to modulate intestinal microbiota and reduce the risk of STm-infected (Buddhasiri et al., 2021; Zhang et al., 2022). This is consistent with our research results between the B. velezensis HBXN2020 treatment group and STm-infected mice. Here, we found that oral B. velezensis HBXN2020 spores raised the colonic microbiota’s alpha diversity and modulated the composition of the intestinal microbiota, namely, enriching the relative abundance of Lactobacillus (known for enhancing the epithelial barrier by increasing mucus secretion and upregulating the expression of tight junction proteins such as claudin-1, occludin, and ZO-1) (Kaur et al., 2021) and Akkermansia (known for modulating intestinal immune response by producing short chain fatty acids and reducing the secretion of proinflammatory cytokines) (Cani et al., 2022), reducing the relative abundance of Enterococcus, Salmonella and Bacteroides (harmful to intestinal homeostasis) (Wang et al., 2023). Notably, we also found that the relative abundance of Lactobacillus and Akkermansia in the ciprofloxacin treatment group was significantly lower than that in the B. velezensis HBXN2020 spores treatment group, indicating that excessive use of antibiotics may interfere with the recovery of beneficial bacteria in the gut. Collectively, infection with S. Typhimurium disrupts gut microbiota and the gut barrier, leading to intestinal inflammation, while oral B. velezensis HBXN2020 spores protects the colon and reduce infection by increasing the abundance of beneficial bacteria and and barrier integrity and decreasing inflammation.

Removed in the revised version: Subsequently, we conducted additional analysis on the colon microbiome. Notably, oral B. velezensis HBXN2020 significantly raised the colonic microbiota’s alpha diversity. As reported in prior studies, elevated alpha diversity is associated with enhanced anti-interference ability (Chen et al., 2021). PCA analysis results showed that oral administration of B. velezensis HBXN2020 significantly affected the structure of the intestinal microbiota. Further analysis at the phylum/genus level revealed that B. velezensis HBXN2020 treatment reduced the increase of Bacteroides, Escherichia-Shigella, and Salmonella (harmful to intestinal homeostasis) in colitis (Wang et al., 2023) but enhanced the relative abundance of Lactobacillus (known for enhancing the epithelial barrier by increasing mucus secretion and upregulating the expression of tight junction proteins such as claudin-1, occludin, and ZO-1) (Kaur et al., 2021) and Akkermansia (known for modulating intestinal immune response by producing short chain fatty acids and reducing the secretion of proinflammatory cytokines) (Cani et al., 2022). Bacteroides are generally ‘friendly’ commensals in the intestine, but they can enter normal tissues through the intestinal mucosa when intestinal barrier function becomes compromised (Polk and Kasper, 1977), leading to inflammation, diarrhea, and abdominal abscesses (Pickard et al., 2017). Escherichia-Shigella is a common intestinal pathogen, and its relative abundance is associated with proinflammatory states (Morgan, 2013). Salmonella can secrete bacterial virulence proteins, which are directly delivered to the host cell cytosol through the Type III secretion system and induce the expression of pro-inflammatory cytokines (Hardt et al., 1998), leading to inflammation. Overall, the S. Typhimurium ATCC14028 infection altered the structure of the colon microbiota and raised the relative abundance of pathogenic bacteria, whereas the therapeutic B. velezensis HBXN2020 increased the abundance of beneficial bacteria, decreased inflammation, and 21 enhanced the function of the intestinal barrier.

Earlier version: We assessed pretreatment in order to learn more about the potential of B. velezensis HBXN2020 spores for disease prevention, given the increased efficacy of treating colitis caused by S. Typhimurium. Similarly to the therapeutic effect of oral B. velezensis HBXN2020, pretreatment with B. velezensis HBXN2020 spores effectively alleviated the symptoms of S. Typhimurium-induced colitis and reduced the expression levels of proinflammatory cytokines and increased anti-inflammatory cytokine levels in the colon of S. Typhimurium-infected mice. Furthermore, pretreatment with B. velezensis HBXN2020 spores also promoted the expression levels of intestinal barrier proteins in the colon of S. Typhimurium-infected mice. In addition, the microbiota composition was also dramatically modulated by prophylactic B. velezensis HBXN2020 in mice with colitis, although the community richness and diversity revealed by the alpha diversity index were not significantly impacted in mice with colitis by prophylactic treatment. Prophylactic B. velezensis HBXN2020 significantly boosted SCFA-producing bacteria, including Lactobacillus and Akkermansia, according to additional genus-level analysis, which was consistent with the effects of therapeutic B. velezensis HBXN2020.

Revised version: We further assessed pretreatment in order to learn more about the potential of B. velezensis HBXN2020 spores for disease prevention, given the increased efficacy of treating infection caused by S. Typhimurium. Similarly to the therapeutic effect of oral B. velezensis HBXN2020, prophylactic B. velezensis HBXN2020 spores effectively alleviated the symptoms of STm-infected mice, namely, reduced the mRNA levels of proinflammatory cytokines and increased the anti-inflammatory cytokine level in the colon of STm-infected mice. Prophylactic B. velezensis HBXN2020 spores also increased the mRNA levels of gut barrier proteins in the colon of STm-infected mice. Furthermore, the microbiota composition was also dramatically modulated by prophylactic B. velezensis HBXN2020 spores in both healthy mice and STm-infected mice. Also, the community richness revealed by the alpha diversity index were not significantly impacted in STm-infected mice by prophylactic B. velezensis HBXN2020 spores. Interestingly, we found that prophylactic B. velezensis HBXN2020 spores signifcantly increased the community diversity in healthy mice, suggesting the potentially benefcial efects of B. velezensis HBXN2020 in gut microenvironment for animals in the future. Moreover, prophylactic B. velezensis HBXN2020 spores significantly boosted SCFA-producing bacteria including Lactobacillus and Akkermansia, which was consistent with the effects of therapeutic oral B. velezensis HBXN2020 spores. Meanwhile, harmful bacteria such as Alistipes, Bacteroides, Escherichia-Shigella, and Enterococcus were significantly enriched in STm-infected mice, reminding us that increased harmful bacteria might play a major role in promoting the development of STm-infected. In summary, our results demonstrated that B. velezensis HBXN2020-mediated microbial communities, especially Lactobacillus and Akkermansia might play a crucial role in alleviating the development of STm-infected.

Removed in the revised version: Similarly, harmful bacteria such as Alistipes, Bacteroides, Escherichia-Shigella, and Enterococcus were significantly enriched in the PBS+STm group. A facultative-pathogenic bacterium called Alistipes encourages the growth of inflammation, obesity, and colorectal cancer (Parker et al., 2020). Based on those results, we found that prophylactic B. velezensis HBXN2020 alleviates S. Typhimurium-induced colitis by increasing the abundance of beneficial bacteria, attenuating proinflammatory responses of the intestine, and increasing intestinal barrier function.

Conclusion

Earlier version: In conclusion, gut mucosal barrier disruption and intestinal inflammation results from dysbiosis of the colonic microbiota caused by S. Typhimurium ATCC14028 infection. Supplementing Bacillus velezensis HBXN2020 can improve intestinal microbiota stability and gut barrier integrity and reduce inflammation to help prevent or treat bacterial colitis.

Revised version: In conclusion, S. Typhimurium ATCC14028 infection can result in gut barrier disruption and intestinal inflammation and intestinal microbiota dysbiosis. Supplementing B. velezensis HBXN2020 spores can improve intestinal microbiota stability and gut barrier integrity and reduce inflammation to help prevent or treat infection by S. Typhimurium. Thus, our results indicate that supplementing B. elezensis HBXN2020 may represent a novel option for preventing Salmonella infections.

Experimental procedures

Added in the revised version: Sample collection and strain isolation

Added in the revised version: Fresh fecal samples were collected from the farms of various locations in southern China including Hubei, Anhui, Hunan, Jiangxi, Guizhou, and Guangxi. All samples were placed in sterile plastic bags, sealed and placed on ice, and immediately transported to the laboratory. Bacillus was isolated following a previously described method with slight modification (Unban et al., 2020). The obtained strains were kept in Luria-Bertani (LB) media containing 25% (v/v) glycerol and stored at −80°C, and prepared in LB medium or LB-Agar medium before use.

Added in the revised version: Antibacterial activity of Bacillus

Added in the revised version: The antibacterial activity of Bacillus isolates was evaluated by using the spot-on-plate method on LB agar plates supplemented with indicator bacteria (Salmonella, E. coli, S. aureus). Briefly, an overnight culture of indicator bacteria was dipped by sterile cotton swabs and spread on LB agar plates. Bacillus solution was spotted onto the double-layer agar plates and incubated at 37°C for 12 h to measure the inhibition zone. A transparent zone of at least 1 mm around the spot was considered positive.

Added in the revised version: The antibacterial spectrum of Bacillus isolates was compared through the agar diffusion test, based on the results of the spot-on-plate test. In brief, an overnight culture of indicator strains was mixed with 10 mL of TSB soft agar (TSB broth containing 0.7% [wt/vol] agar) and poured into a sterile plate covered with LB agar and Oxford cups, 8 mm diameter wells were prepared in the TSB agar after removing the cups. The wells were filled with 100 μL of Bacillus solution. The plates were incubated at 37°C for 14 h, and the inhibition zone was measured.

Removed in the revised version: Animals and bacterial strains

Removed in the revised version: Female specific-pathogen-free (SPF) C57BL/6 mice aged 6 to 8 weeks were purchased from the Animal Experimental Center of Huazhong Agricultural University. All mice experiments were conducted in the standard SPF facility of Huazhong Agricultural University, with 12 h of light and 12 h of darkness at a temperature of 25°C and ad libitum access to food and water. The use of animals in this experiment was approved by the Animal Care and Ethics Committee of Huazhong Agricultural University (Ethics Approval Number: HZAUMO-2023-0089).

Earlier version: The strains used in this study are listed in Table S3, and all the primers used in Table S4. Clostridium perfringens were cultured in a fluid thioglycollate medium (FTG) (Hopebio, Qingdao, China) at 45°C. Streptococcus suis and Pasteurella multocida were cultured in tryptic soy broth (TSB) (BD Biosciences, MD, United States) medium supplemented with 5% (v/v) sheep serum (Solarbio, Beijing, China) at 37°C, and Actinobacillus pleuropneumoniae were cultured in TSB supplemented with 5% (v/v) sheep serum and 5mM nicotinamide adenine dinucleotide (NAD) at 37°C, while the other bacterial strains were cultured in Luria Bertani (LB) medium (Solarbio, Beijing, China) at 37°C. In addition, the Difco sporulation medium (DSM) was used for inducing the sporulation of B. velezensis HBXN2020 via the nutrient depletion method (Tang et al., 2017).

Revised version: The indicator strains used in this study are listed in Table S4, and all the primers used in Table S5. Clostridium perfringens were cultured in a fluid thioglycollate medium (FTG) (Hopebio, Qingdao, China) at 45°C. Streptococcus suis and Pasteurella multocida were cultured in TSB medium supplemented with 5% (v/v) sheep serum (Solarbio, Beijing, China) at 37°C, and Actinobacillus pleuropneumoniae were cultured in TSB supplemented with 5% (v/v) sheep serum and 5mM nicotinamide adenine dinucleotide (NAD) at 37°C, while the other bacterial strains were cultured in Luria Bertani (LB) medium (Solarbio, Beijing, China) at 37°C. In addition, the Difco sporulation medium (DSM) was used for inducing the sporulation of Bacillus via the nutrient depletion method (Tang et al., 2017).

Removed in the revised version: Stain isolation, growth conditions and strain identification

Removed in the revised version: In this study, we isolated 362 Bacillus strains from the feces of healthy piglets, and four Bacillus strains were screened through initial antibacterial tests. Then, a strain with broad-spectrum antimicrobial activity was ultimately screened out by the antibacterial spectrum assay (Table S5), and named as HBXN2020. The strain was stored in 28% glycerol at −80°C. For culturing the bacterial cells, Luria-Bertani (LB) media or agar plates supplemented with 1.5% agar were used and incubated at 37°C.

Removed in the revised version: A single bacterial colony was inoculated into 10 mL of LB medium and incubated overnight at 37°C with shaking (180 rpm). The bacteria genomic DNA was extracted using the TIANamp Genomic DNA Kit (TIANGEN, Beijing, China) from bacteria. Then, 16S ribosomal sequence was amplified using 16S-specific primers and sequenced to identify the HBXN2020.

Removed in the revised version: Genome sequencing, annotation and analysis

Removed in the revised version: The complete genome of B. velezensis HBXN2020 was sequenced using the Illumina HiSeq PE and PacBio RSII (SMRT) platforms (Shanghai Majorbio Bio-pharm Technology Co., Ltd.). Briefly, the short reads from the Illumina HiSeq PE were assembled into contigs using SPAdenovo (http://soap.genomics.org.cn/). The long reads from the PacBio RSII and the Illumina contigs were then aligned using the miniasm and Racon tools in Unicycler (version 0.4.8, https://github.com/rrwick/Unicycler) to generate long-read sequences. During the assembly process, sequence correction was performed using Pilon (version 1.22, https://github.com/broadinstitute/pilon/wiki/Standard-Output). Lastly, a complete genome with seamless chromosomes was obtained.

Removed in the revised version: The coding sequences (CDs) of the B. velezensis HBXN2020 genome were predicted using Glimmer (version 3.02, http://ccb.jhu.edu/software/glimmer/index.shtml). The tRNA and rRNA genes were predicted using TRNAscan-SE (version 2.0, http://trna.ucsc.edu/software/) and Barrnap (version 0.8, https://github.com/tseemann/barrnap), respectively. The functional annotation of all CDSs was performed using various databases, including Swiss-Prot Database (https://web.expasy.org/docs/swiss-prot_guideline.html), Pfam Database (http://pfam.xfam.org/), EggNOG Database (http://eggnog.embl.de/), GO Database (http://www.geneontology.org/), and KEGG Database (http://www.genome.jp/kegg/). The circular map of the B. velezensis HBXN2020 genome was generated using CGView (version 2, http://wishart.biology.ualberta.ca/cgview/) (Stothard and Wishart, 2005).

Removed in the revised version: Comparative genomic analysis

Removed in the revised version: To elucidate the phylogenetic relationships from a whole-genome perspective, a phylogenetic tree based on the whole genome was constructed using the Type (Strain) Genome Server (TYGS) online (https://ggdc.dsmz.de/) (Meier-Kolthoff et al., 2022). The average nucleotide identity (ANI) values of genome to genome were calculated using the JSpeciesWS online service (Richter et al., 2016), and a heatmap was then generated using TBtools (version 1.113) (Chen et al., 2020).

Earlier version: Growth curves of B. velezensis HBXN2020

Revised version: Growth curves of HBXN2020

Earlier version: The growth curve of B. velezensis HBXN2020 was recorded in flat-bottomed 100-well microtiter plates via detecting optical density at 600 nm (OD600) at 1 h intervals using the automatic growth curve analyzer (Bioscreen, Helsinki, Finland).

Revised version: The growth curve of HBXN2020 was recorded in flat-bottomed 100-well microtiter plates via detecting optical density at 600 nm (OD600) at 1 h intervals using the automatic growth curve analyzer (Bioscreen, Helsinki, Finland).

Added in the revised version: In vitro resistance assay of HBXN2020

Added in the revised version: HBXN2020 spores (100 μL) or vegetative cells (100 μL) were separately resuspended in 900 μL of LB medium supplemented with different pH values (2, 3, 4, 5 or 6), bile salts (0.85% NaCl, 0.3%), simulated gastric fluid (SGF, HCl, pH 1.2) containing 10 g L-1 of pepsin in 0.85% NaCl solution, or simulated intestinal fluid (SIF, NaOH, pH 6.8) containing 10 g L-1 of trypsin in 0.05 M KH2PO4 solution, and incubated at 37 °C. A normal LB medium (pH 7.0) was used as the control. At predetermined time points, 100 μL was taken from each sample, serially diluted 10-fold with sterile PBS (pH 7.2), and then spread onto LB agar plates. The plates were incubated overnight in a constant temperature incubator at 37°C, and the bacterial colonies were counted. The survival rate was calculated using the following formula: Survival rate = (number of bacteria in the treatment group/number of bacteria in the control group) × 100%.

Added in the revised version: One milliliter of HBXN2020 spores or vegetative cells were separately placed in water baths at different temperatures (37°C, 45°C, 55°C, 65°C, 75°C, 85°C or 95°C) for 20 min, with a 37°C water bath used as the control. The survival rate was calculated as described above.

Removed in the revised version: Antimicrobial Assays

Removed in the revised version: The in vitro antagonistic activity of B. velezensis HBXN2020-CFS was tested using the agar well-diffusion method against 18 indicator strains (pathogens), which included 7 standard strains (E. coli ATCC 25922, E. coli ATCC 35150, S. Typhimurium ATCC 14028 (STm), S. Typhimurium SL1344, S. aureus ATCC 29213, S. aureus ATCC 43300, and C. perfringens CVCC 2030) and 11 clinical isolates (E. coli EC024, S. Enteritidis SE006, S. aureus S21, C. perfringens CP023, C. perfringens CP002, S. suis SC19, S. suis SS006, P. multocida PM002, P. multocida PM008, A. pleuropneumoniae APP015, and A. pleuropneumoniae APP017). All plates were cultured at 37°C for 16 h before observing the inhibition zone, and the diameter of the inhibition zone was measured using a vernier caliper. The presence of a clear zone indicated antagonistic activity.

Antibiotic susceptibility assays

Antimicrobial susceptibility testing was performed using the Kirby-Bauer (KB) disk diffusion method in accordance with the Clinical Laboratory Standards Institute (CLSI) guidelines (CLSI., 2018.). The antimicrobial agents tested were: Ampicillin, Meropenem, Piperacillin, Gentamycin, Tetracycline, Doxycycline, Minocycline, Erythromycin, Enrofloxacin, Ofloxacin, Sulfamethoxazole, Trimethoprim-sulfamethoxazole, Polymyxin B, Teicoplanin, Trimethoprim, Florfenicol, Spectinomycin, Nitrofurantoin, and Rifampicin. The diameter of the inhibition zone was measured using a vernier caliper.

Added in the revised version: Antimicrobial Assays

Added in the revised version: The in vitro antagonistic activity of HBXN2020-CFS was tested using the agar well-diffusion method against 18 indicator strains (pathogens), which included 7 standard strains (E. coli ATCC 25922, E. coli ATCC 35150, S. Typhimurium ATCC 14028 (STm), S. Typhimurium SL1344, S. aureus ATCC 29213, S. aureus ATCC 43300, and C. perfringens CVCC 2030) and 11 clinical isolates (E. coli EC024, S. Enteritidis SE006, S. aureus S21, C. perfringens CP023, C. perfringens CP002, S. suis SC19, S. suis SS006, P. multocida PM002, P. multocida PM008, A. pleuropneumoniae APP015, and A. pleuropneumoniae APP017). All plates were cultured at 37°C for 16 h before observing the inhibition zone, and the diameter of the inhibition zone was measured using a vernier caliper. The presence of a clear zone indicated antagonistic activity.

Removed in the revised version: In vitro resistance assay of B. velezensis HBXN2020

Removed in the revised version: B. velezensis HBXN2020 spores (100 μL) or vegetative cells (100 μL) were separately resuspended in 900 μL of LB medium supplemented with different pH values (2, 3, 4, 5 or 6), bile salts (0.85% NaCl, 0.3%), simulated gastric fluid (SGF, HCl, pH 1.2) containing 10 g L-1 of pepsin in 0.85% NaCl solution, or simulated intestinal fluid (SIF, NaOH, pH 6.8) containing 10 g L-1 of trypsin in 0.05 M KH2PO4 solution, and incubated at 37 °C. A normal LB medium (pH 7.0) was used as the control. At predetermined time points, 100 μL was taken from each sample, serially diluted 10-fold with sterile PBS (pH 7.2), and then spread onto LB agar plates. The plates were incubated overnight in a constant temperature incubator at 37°C, and the bacterial colonies were counted. The survival rate was calculated using the following formula: Survival rate = (number of bacteria in the treatment group/number of bacteria in the control group) × 100%.

Removed in the revised version: One milliliter of B. velezensis HBXN2020 spores or vegetative cells were separately placed in water baths at different temperatures (37°C, 45°C, 55°C, 65°C, 75°C, 85°C or 95°C) for 20 min, with a 37°C water bath used as the control. The survival rate was calculated as described above.

In vitro bacterial competition

Earlier version: To investigate whether B. velezensis HBXN2020 directly inhibited the growth of STm, we performed spot-on lawn and agar-well diffusion assays, as well as co-culture assays in liquid culture medium. In the spot-on lawn antimicrobial assays, we prepared double layers of agar by first pouring LB agar into the plate as the bottom layer. The top layer consisted of 10 mL TSB broth containing 0.7% agar with STm overnight culture. Then, 10 μL of B. velezensis HBXN2020 overnight culture and cell-free supernatant (CFS) were respectively spotted onto TSB agar and incubated at 37°C for 12 h to measure the inhibition zone. A transparent zone of at least 1 mm around the spot was considered positive.

Revised version: To investigate whether HBXN2020 directly inhibited the growth of STm, we performed spot-on lawn and agar-well diffusion assays, as well as co-culture assays in liquid culture medium. In the spot-on lawn antimicrobial assays, we prepared double layers of agar by first pouring LB agar into the plate as the bottom layer. The top layer consisted of 10 mL TSB broth containing 0.7% agar with STm overnight culture. Then, 10 μL of HBXN2020 overnight culture and cell-free supernatant (CFS) were respectively spotted onto TSB agar and incubated at 37°C for 12 h to measure the inhibition zone. A transparent zone of at least 1 mm around the spot was considered positive.

Earlier version: The antagonistic effect of B. velezensis HBXN2020-CFS against STm was determined using the agar-well diffusion assays. To collect B. velezensis HBXN2020-CFS, the culture was centrifuged at 9,000g for 15 min at 4°C, and the supernatant was filtered through a 0.22 μm membrane filter (Millipore, USA). The STm lawn medium was prepared by mixing 10 mL of TSB broth containing 0.7% agar with STm overnight culture and then poured into a sterile plate covered with LB agar and Oxford cups, 8 mm diameter wells were prepared in the TSB agar after removing the cups. The wells were filled with 100 μL of B. velezensis HBXN2020-CFS, LB medium or ampicillin (100 µg mL-1). The plates were incubated at 37°C for 14 h, and the inhibition zone was measured.

Revised version: The antagonistic effect of HBXN2020-CFS against STm was determined using the agar-well diffusion assays. To collect HBXN2020-CFS, the culture was centrifuged at 9,000g for 15 min at 4°C, and the supernatant was filtered through a 0.22 μm membrane filter (Millipore, USA). The STm lawn medium was prepared by mixing 10 mL of TSB broth containing 0.7% agar with STm overnight culture and then poured into a sterile plate covered with LB agar and Oxford cups, 8 mm diameter wells were prepared in the TSB agar after removing the cups. The wells were filled with 100 μL of HBXN2020-CFS, LB medium or ampicillin (100 µg mL-1). The plates were incubated at 37°C for 14 h, and the inhibition zone was measured.

Earlier version: The co-culture assay was conducted by incubating B. velezensis HBXN2020 and S. Typhimurium ATCC14028 (carry pET28a (+), kanamycin resistance) (Cao et al., 2019) separately overnight at 37°C and diluting them to 104 CFU/mL. Then, the two strains were mixed at different ratios (1:1, 1:10, 1:50 or 1:100) and co-cultured at 37°C with shaking (180 rpm). At predetermined time points, serial 10-fold dilutions were prepared for all samples and spread onto selective (kanamycin 50 µg mL-1) LB agar plates and cultured at 37°C for 12 h before bacterial counting. Viable colony counts ranged from 30 to 300 per plate.

Revised version: The co-culture assay was conducted by incubating HBXN2020 and S. Typhimurium ATCC14028 (carry pET28a (+), kanamycin resistance) (Cao et al., 2019) separately overnight at 37°C and diluting them to 104 CFU/mL. Then, the two strains were mixed at different ratios (1:1, 1:10, 1:50 or 1:100) and co-cultured at 37°C with shaking (180 rpm). At predetermined time points, serial 10-fold dilutions were prepared for all samples and spread onto selective (kanamycin 50 µg mL-1) LB agar plates and cultured at 37°C for 12[h before bacterial counting. Viable colony counts ranged from 30 to 300 per plate.

Added in the revised version: Genome sequencing, annotation and analysis

Added in the revised version: The complete genome of strain was sequenced using the Illumina HiSeq PE and PacBio RSII (SMRT) platforms (Shanghai Majorbio Bio-pharm Technology Co., Ltd.). Briefly, the short reads from the Illumina HiSeq PE were assembled into contigs using SPAdenovo (http://soap.genomics.org.cn/). The long reads from the PacBio RSII and the Illumina contigs were then aligned using the miniasm and Racon tools in Unicycler (version 0.4.8, https://github.com/rrwick/Unicycler) to generate long-read sequences. During the assembly process, sequence correction was performed using Pilon (version 1.22, https://github.com/broadinstitute/pilon/wiki/Standard-Output). Lastly, a complete genome with seamless chromosomes was obtained.

Added in the revised version: The coding sequences (CDs) of the strain genome were predicted using Glimmer (version 3.02, http://ccb.jhu.edu/software/glimmer/index.shtml). The tRNA and rRNA genes were predicted using TRNAscan-SE (version 2.0, http://trna.ucsc.edu/software/) and Barrnap (version 0.8, https://github.com/tseemann/barrnap), respectively. The functional annotation of all CDSs was performed using various databases, including Swiss-Prot Database (https://web.expasy.org/docs/swiss-prot_guideline.html), Pfam Database (http://pfam.xfam.org/), EggNOG Database (http://eggnog.embl.de/), GO Database (http://www.geneontology.org/), and KEGG Database (http://www.genome.jp/kegg/). The circular map of the strain genome was generated using CGView (version 2, http://wishart.biology.ualberta.ca/cgview/) (Stothard and Wishart, 2005).

Added in the revised version: Comparative genomic analysis

Added in the revised version: To elucidate the phylogenetic relationships of the strain from a whole-genome perspective, a phylogenetic tree based on the whole genome was constructed using the Type (Strain) Genome Server (TYGS) online (https://ggdc.dsmz.de/) (Meier-Kolthoff et al., 2022). The average nucleotide identity (ANI) values of the strain genome to the reference strain genome were then calculated using the JSpeciesWS online service (Richter et al., 2016), and a heatmap was generated using TBtools (version 1.113) (Chen et al., 2020).

Added in the revised version: In vivo safety evaluation of B. velezensis HBXN2020

Added in the revised version: Female specific-pathogen-free (SPF) C57BL/6 mice aged 6 to 8 weeks were purchased from the Animal Experimental Center of Huazhong Agricultural University. All mice experiments were conducted in the standard SPF facility of Huazhong Agricultural University, with 12 h of light and 12 h of darkness at a temperature of 25°C and ad libitum access to food and water. The use of animals in this experiment was approved by the Animal Care and Ethics Committee of Huazhong Agricultural University (Ethics Approval Number: HZAUMO-2023-0089).

Removed in the revised version: Biosafety assessment

Earlier version: Biosafety assay of B. velezensis HBXN2020 referred to the method described by Zhou et al.(Zhou et al., 2022) with slight modifications. After a 7-day acclimation period (free access to water and food), mice were randomly divided into four treatment groups (n = 5): low-dose B. velezensis HBXN2020 spores group (L-HBXN2020 group, 107 CFU/mouse), medium-dose B. velezensis HBXN2020 spores group (M-HBXN2020 group, 108 CFU/mouse), high-dose B. velezensis HBXN2020 spores group (H-HBXN2020 group, 109 CFU/mouse), and a control group. During the experimental period (15 days), mice were weighed and orally gavaged with their respective treatments once every two days. On day 15, all mice were euthanized, and blood, heart, liver, spleen, lung, kidney, ileum, cecum, and colon were collected. Blood samples were used for routine blood and biochemistry tests. Major organ tissues (heart, liver, spleen, lung, and kidney) and a 5-mm distal segment of the colon were used for histopathology. The remaining colonic tissues were rapidly frozen in liquid nitrogen and stored at −80°C for cytokine and tight junction protein expression analysis.

Revised version: Safety assay of B. velezensis HBXN2020 referred to the method described by Zhou et al. (Zhou et al., 2022) with slight modifications. After a 7-day acclimation period (free access to water and food), mice were randomly divided into four treatment groups (n = 5): low-dose B. velezensis HBXN2020 spores group (L-HBXN2020 group, 107 CFU/mouse), medium-dose B. velezensis HBXN2020 spores group (M-HBXN2020 group, 108 CFU/mouse), high-dose B. velezensis HBXN2020 spores group (H-HBXN2020 group, 109 CFU/mouse), and a control group. During the experimental period (15 days), mice were weighed and orally gavaged with their respective treatments once every two days. On day 15, all mice were euthanized, and blood, heart, liver, spleen, lung, kidney, ileum, cecum, and colon were collected. Blood samples were used for routine blood and biochemistry tests. Major organ tissues (heart, liver, spleen, lung, and kidney) and a 5-mm distal segment of the colon were used for histopathology. The remaining colonic tissues were rapidly frozen in liquid nitrogen and stored at −80°C for cytokine and tight junction protein expression analysis.

Earlier version: Salmonella Typhimurium-induced mouse model of colitis

Revised version: Salmonella Typhimurium-infected mouse model

Earlier version: After 7 days of acclimation (free access to water and food), mice were randomly divided into three treatment groups (n = 8): Control group, STm + PBS group, and STm + HBXN2020 group. The S. Typhimurium-induced colitis was performed as previously described (Stecher et al., 2005), with slight modifications. On the first day of the experiment, all mice in the STm + PBS group and STm + HBXN2020 group were orally inoculated with 200 μL (5×107 CFU/mouse) of STm. On days 1, 3 and 5 following STm infection, each mouse in the STm + HBXN2020 group received 200 μL (1×108 CFU/mouse) of B. velezensis HBXN2020 spores via gavage administration. In contrast, the control group and STm + PBS group received 200 μL of sterile PBS by oral gavage. Fecal samples were collected daily following STm infection and resuspended in sterile PBS. The number of STm in mice feces from both the STm + PBS and STm + HBXN2020 groups was then determined by spreading a serial 10-fold dilution on selective LB agar plates containing 50 µg mL-1 kanamycin. Throughout the entire experiment, the body weight, stool consistency and fecal occult blood of all mice were monitored daily. As shown in Table S6, disease activity index (DAI) was calculated by the sum of the scores from three parameters (Praveschotinunt et al., 2019). On day 7 after STm infection, all mice were euthanized, their ileum, cecum and colon were collected. The length of colon was measured, and a 5-mm distal segment of the colon was fixed in 4% paraformaldehyde for further histopathology. The remaining colon was then rapidly frozen in liquid nitrogen and stored at −80°C for cytokine and tight junction protein expression analysis. Then, the number of STm in the ileum, cecum, and colon was determined by spreading serial 10-fold dilutions on selective LB agar plates.

Revised version: After 7 days of acclimation (free access to water and food), mice were randomly divided into five treatment groups (n = 6): control group, HBXN2020 group, STm+PBS group, STm+HBXN2020 group and STm+CIP group. The S. Typhimurium-infected model was performed as previously described (Stecher et al., 2005), with slight modifications. On the first day of the experiment, all mice in the STm+PBS group, STm+HBXN2020 group and STm+CIP group were orally inoculated with 200 μL (5×107 CFU/mouse) of STm. On days 1, 3 and 5 following STm infection, each mouse in the HBXN2020 group, STm+HBXN2020 group and STm+CIP group received 200 μL (1×108 CFU/mouse) of B. velezensis HBXN2020 spores or ciprofloxacin (1 mg/mL) via gavage administration, respectively. In contrast, the control group and STm+PBS group received 200 μL of sterile PBS by oral gavage. Fecal samples were collected daily following STm infection and resuspended in sterile PBS. The number of STm in mice feces from both the STm+PBS, STm+HBXN2020 and STm+CIP groups was then determined by spreading a serial 10-fold dilution on selective LB agar plates containing 50 µg mL-1 kanamycin. Throughout the entire experiment, the body weight, stool consistency and fecal occult blood of all mice were monitored daily. As shown in Table S6, disease activity index (DAI) was calculated by the sum of the scores from three parameters (Praveschotinunt et al., 2019). On day 7 after STm infection, all mice were euthanized, their ileum, cecum and colon were collected. The length of colon was measured, and a 5-mm distal segment of the colon was fixed in 4% paraformaldehyde for further histopathology. The remaining colon was then rapidly frozen in liquid nitrogen and stored at −80°C for cytokine and tight junction protein expression analysis. Then, the number of STm in the ileum, cecum, and colon was determined by spreading serial 10-fold dilutions on selective LB agar plates.

Earlier version: To investigate the prophylactic efficacy of B. velezensis HBXN2020 in ameliorating STm-induced colitis, another independent experiment was conducted using six-week-old female C57BL/6 mice (SPF). After a seven-day acclimation period, mice were randomly assigned to three groups (n = 8): control group, PBS + STm group and HBXN2020 + STm group. On days 1, 3, 5 and 7, each mouse in the HBXN2020 + STm group received 200 μL (1×108 CFU/mouse) of B. velezensis HBXN2020 spores via gavage administration. Meanwhile, mice in the control group and PBS + STm group received 200 μL of sterile PBS by oral gavage. On day 7, all mice in the PBS + STm group and HBXN2020 + STm group were orally inoculated with 200 μL (5×107 CFU/mouse) of STm. Fecal samples were collected daily following STm infection from both groups and resuspended in sterile PBS. The number of STm in the feces of mice was then determined by spreading serial 10-fold dilutions on selective LB agar plates (50 µg mL-1 kanamycin). Throughout the entire experiment, the body weight and DAI scores of all mice were monitored daily. On day 12, all mice were euthanized, and the ileum, cecum, and colon were collected. The colon length was measured, and a 5-mm distal segment of the colon was fixed in 4% formalin for sectioning and staining. The remaining colon was stored at −80°C for future analysis. Lastly, the number of STm in the ileum, cecum, and colon was determined using a selective LB agar plate.

Revised version: To investigate the prophylactic efficacy of B. velezensis HBXN2020 in ameliorating STm-infected, another independent experiment was conducted using six-week-old female C57BL/6 mice (SPF). After a seven-day acclimation period, mice were randomly assigned to five groups (n = 6): control group, HBXN2020 group, PBS+STm group, HBXN2020+STm group and CIP+STm group. On days 1, 3, 5 and 7, each mouse in the HBXN2020 group, HBXN2020+STm group and CIP+STm group received 200 μL (1×108 CFU/mouse) of B. velezensis HBXN2020 spores or ciprofloxacin (1 mg/mL) via gavage administration. Meanwhile, mice in the control group and PBS+STm group received 200 μL of sterile PBS by oral gavage. On day 7, all mice in the PBS+STm group, HBXN2020+STm group and CIP+STm group were orally inoculated with 200 μL (5×107 CFU/mouse) of STm. Fecal samples were collected daily following STm infection from both groups and resuspended in sterile PBS. The number of STm in the feces of mice was then determined by spreading serial 10-fold dilutions on selective LB agar plates (50 µg mL-1 kanamycin). Throughout the entire experiment, the body weight and DAI scores of all mice were monitored daily. On day 12, all mice were euthanized, and the ileum, cecum, and colon were collected. The colon length was measured, and a 5-mm distal segment of the colon was fixed in 4% formalin for sectioning and staining. The remaining colon was stored at −80°C for future analysis. Lastly, the number of STm in the ileum, cecum, and colon was determined using a selective LB agar plate.

Determination of cytokines and tight junction protein expression in colon tissue

Earlier version: Total RNA was extracted from colon tissues using TRIpure reagent (Aidlab, China), and cDNA was obtained using HiScript® III RT SuperMix for qPCR (+gDNA wiper) (Vazyme, China). RT-qPCR for each gene was performed in triplicate using qPCR SYBR Green Master Mix (Yeasen Biotechnology, Shanghai, China). The relative expression level of cytokine and tight junction protein genes was calculated using the 2−ΔΔCt method with β-actin and GAPDH as reference genes. The primer sequences used in the RT-qPCR test are listed in Table S4.

Revised version: Total RNA was extracted from colon tissues using TRIpure reagent (Aidlab, China), and cDNA was obtained using HiScript® III RT SuperMix for qPCR (+gDNA wiper) (Vazyme, China). RT-qPCR for each gene was performed in triplicate using qPCR SYBR Green Master Mix (Yeasen Biotechnology, Shanghai, China). The relative expression level of cytokine and tight junction protein genes was calculated using the 2−ΔΔCt method with β-actin and GAPDH as reference genes. The primer sequences used in the RT-qPCR test are listed in Table S5.

Histopathology analysis

Colon tissue samples (0.5 cm) were fixed in 4% paraformaldehyde for 24 h, and the fixed tissues were then embedded in paraffin and sectioned. The sections were stained with hematoxylin and eosin (H&E) and observed and imaged using an optical microscope (Olympus Optical, Tokyo, Japan). The histopathological score included the degree of inflammatory infiltration, changes in crypt structures, and the presence or absence of ulceration and edema. The scoring criteria were determined as previously described (Wu et al., 2022).

16S rRNA gene sequencing and analysis

According to the manufacturer’s instructions, colon microbial community genomes were extracted using E.Z-N.A ® Stool DNA Kit (Omega; D4015-01), and quality was detected by 1% agarose gel electrophoresis. The 16S rRNA V3-V4 variable region was amplified by PCR using universal primers 338F (5 ’- ACTCCTACGGGGGGCAG-3’) and 806R (5 ’- GACTACHVGGGTWTCTAAT-3’). The PCR products were examined by electrophoresis on 2% agarose gels and then purified with the AxyPrep DNA gel extraction kit (Axygen Biosciences, USA). A sequencing library was constructed using the NEXTFLEX® Rapid DNA-Seq Kit, and sequencing was performed using the Illumina MiSeq platform. Raw reads were quality evaluated and filtered by fastp (version 0.20.0) and merged using FLASH (version 1.2.7). The optimized sequences were clustered into operational taxonomic units (OTUs) based on 97% sequence similarity using UPARSE (version 7.1). The representative sequences of each OTU was classified by RDP classifier (version 2.2; confidence threshold value, 0.7). Alpha diversity was assessed using the ACE, chao, Shannon, and Simpson indices. The β- diversity analysis was performed using Bray-Curtis distances and visualized through principal component analysis (PCA). Linear discriminant analysis (LDA) effect size (LefSe) was used to identify differential microbiota between groups.

Statistical Analysis

Statistical analysis was performed using GraphPad Prism 8.3.0 (GraphPad Software, San Diego, CA, USA) and Excel (Microsoft, Redmond, USA). Data are presented as mean ± standard error of the mean (SEM). Differences between two groups were evaluated using two-tailed unpaired Student’s t-test, and all other comparisons were conducted using one-way analysis of variance (ANOVA) with Tukey’s multiple comparisons test. For all analyses, significance differences are denoted as: *, p < 0.05, **, p < 0.01, and ***, p < 0.001.

Removed in the revised version: Data availability

Removed in the revised version: The datasets supporting the conclusions of this article were deposited in the NCBI Sequence Read Archive database under the accession numbers: CP119399.1 (whole genome sequencing raw data) and PRJNA1045315 (microbiota raw sequencing data).

Removed in the revised version: Supporting information